# SLU05 Exercise notebook 0
This exercise notebook is for complete beginners to practice the syntax of the loops and conditions. If they start to feel boring, it's a good sign that you have learned the lesson and you are ready to move to exercise notebook 1. ;)

```python
# For evaluation purposes, please run this cell
import hashlib
import random
random.seed(a=2)
from utils import card_deck
```

## 1. Boolean expressions

### Exercise 1.1 Comparison operators with numbers
Fill in comparison operators so that each expression evaluates to `True`. Example:

What is the comparison operator that I should fill in here?

```python
expression_0 = 10 comparison_operator 1
```

10 is greater than 1, so I will fill in `>`:

```python
expression_0 = 10 > 1
```

To check my solution, I will print `expression_0`:

```python
print(expression_0)
```

Which gives:

```python
True
```




<!-- grade_id: cell-bb14e18f7e6f7895, points: 0, grade: false, solution: true -->
```python
# Uncomment each line and fill in the correct comparison operator

#expression_1 = 3**2 comparison_operator 2**3
#expression_2 = 1+2+3+4 comparison_operator 5+6+7-1-2-3
#expression_3 = 17 * (6+7) comparison_operator 13 * (8+9)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-59fa3641918cf2e1, points: 1, grade: true, solution: false -->
```python
assert expression_1, 'check expression_1'
assert expression_2, 'check expression_2'
assert expression_3, 'check expression_3'
"Well done!"
```

### Exercise 1.2 Comparison operators with strings
In the following expressions we are comparing strings. Unfortunately, all the comparisons evaluate to `False`. Change the strings in the expressions according to the instructions to make the expressions `True`. Example:

Change **one character** in the strings in the following expression:

```python
expression_0s = 'ana' == 'Ana'
```

I see that these strings are not equal because of the first character which is lowercase in one and uppercase in the other. So I change the one of them like this:

```python
expression_0s = 'Ana' == 'Ana'
```

Voilá, now the expression evaluates to `True`!

<!-- grade_id: cell-7a5c491455009d46, points: 0, grade: false, solution: true -->
```python
# Add one character to one of the strings to make the expression True
expression_1s = 'Catch' + 'me!'  == 'Catch me!'

# Change one character in the strings to make the expression True.
expression_2s = 'what a beautiful day!' > 'what a beautiful day!'

# Change one character in the string to make the expression True
expression_3s = '100 + 1'  >= str(100 + 1)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f1b59a731d1b06f4, points: 1, grade: true, solution: false -->
```python
assert expression_1s, 'check expression_1s'
assert expression_2s, 'check expression_2s'
assert expression_3s, 'check expression_3s'
'Well done!'
```

### Exercise 1.3 Boolean expressions
Consider the following expressions. Without running the code, evaluate their truth value and assign the answer to the corresponding answer variable.

```python
boolean_1 = 15 != 1 + 5 or 15 < 1 + 5
boolean_2 = 'cat' in ['fly', 'cat', 'fish'] and 'kitten' > 'dog'
boolean_3 = 'yes' > 'no' and bool(15 % 3)
boolean_4 = 2 < 1 or 2 > 1 and 2!=1 and not 2 == 1
boolean_5 = not 'w' >= 'W' or 'W' =='w' and 'W' <= 'w' 
```


<!-- grade_id: cell-71016edc78ab585b, points: 0, grade: false, solution: true -->
```python
#boolean_1 =
#boolean_2 =
#boolean_3 = 
#boolean_4 = 
#boolean_5 = 

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-7f7b5545409bda29, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha1(bytes(str(boolean_1)+str(boolean_2)+str(boolean_3)+str(boolean_4)+str(boolean_5), encoding='utf8')).hexdigest() == \
'f47a14798a4423871234186c100ea1976ddf0b2a', 'One of the answers is not correct.'
'Well done!'
```

### Exercise 1.4 Boolean expressions 2
Consider the following expressions containing variables. All of them evaluate to `True`. Knowing the truth value of the expressions, what is the truth value of the variables? Do not run the code, just think. :) Assign the answer to the corresponding answer variable.

```python
variable_1 and '10' != '9+1'

variable_2 or 4 * 5 > 5 * 4

1 - 2 * 10 != 21 - 5 * 8 and 'Night' + 'Day' < 'Sun' + 'Stars'  or not variable_3
```

<!-- grade_id: cell-cd973652ed94dcfe, points: 0, grade: false, solution: true -->
```python
#variable_1 = 
#variable_2 = 
#variable_3 = 


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3ba2a592cee6a423, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha1(bytes(str(boolean_1)+str(boolean_2)+str(boolean_3)+str(boolean_4)+str(boolean_5), encoding='utf8')).hexdigest() == \
'f47a14798a4423871234186c100ea1976ddf0b2a', 'One of the answers is not correct.'
'Well done!'
```

## 2. If statements

### Exercise 2.1
Below is a piece of code with an `if` statement whose condition is missing. Choose the correct option for the condition so that the code prints `blue` if the `my_variable` is larger than 10 and `red` otherwise.

```python
my_variable = 1000
if condition:
    print('blue')
else:
    print('red')
```

&emsp;A: my_variable > 10   
&emsp;B: my_variable <= 10   
&emsp;C: my_variable < 10  
&emsp;D: my_variable >= 10

<!-- grade_id: cell-30f2c187da094a99, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

#exercise_2_1 =
```

<!-- grade_id: cell-830a01d0d6639596, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha1(bytes(exercise_2_1, encoding='utf8')).hexdigest() == '6dcd4ce23d88e2ee9568ba546c007c63d9131c1b', "Not the correct option."
```

### Exercise 2.2
Below is a `for` loop that assigns the elements of the list `before` to the list `after` one by one. The `after` list is printed at the end. Run the code if you like to see what it does.

Add an `if` statement inside the `for` loop which filters the elements of the `before` list so that only elements divisible by 2 are assigned to the `after` loop.

**Do not remove the code that is already in the cell!** Just add the `if` statement where indicated and adjust the indenting if necessary.

Just in case the original list goes missing, rerun the cell below to get it again.

```python
# Rerun this cell if you lose the before list
before = [31, 16, 32, 18, 151, 55, -5, 3, 1024]
```

<!-- grade_id: cell-d57af371972d6aae, points: 0, grade: false, solution: true -->
```python
before = [31, 16, 32, 18, 151, 55, -5, 3, 1024]
after = []

for number in before:
    # add the if statement here
    after.append(number)

print(after)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fe83b9cd46566bc9, points: 1, grade: true, solution: false -->
```python
assert before == [31, 16, 32, 18, 151, 55, -5, 3, 1024], 'Do not change the before list!'
assert len(after) == 4, 'The length of the after list is not correct.'
assert sum(after) == 1090, 'The elements of the after list are not correct.'
print(after)
'Well done!'
```

### Exercise 2.3
Below is a for loop that assigns the elements of the list `some_numbers` to the lists `positive` and `other` one by one. Both lists are printed at the end. Run the code if you like to see what it does.

Add an `if-else` statement inside the `for` loop which filters the elements of the `some_numbers` list so that positive numbers are assigned to the `positive` list and all other numbers are assigned to the `other` list.

**Do not remove the code that is already in the cell!** Just add the `if-else` statement where indicated and adjust the indenting if necessary.

```python
# Rerun this cell if you lose the some_numbers list
some_numbers = [-1, 10, 11, 21, -3, -107, 16]
```

<!-- grade_id: cell-df57837406e0c11e, points: 0, grade: false, solution: true -->
```python
some_numbers = [-1, 10, 11, 21, -3, -107, 16]
positive = []
other = []

for number in some_numbers:
    # add the if-else statement here
    positive.append(number)
    other.append(number)

print(positive)
print(other)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1efcba26056fadd3, points: 1, grade: true, solution: false -->
```python
assert some_numbers == [-1, 10, 11, 21, -3, -107, 16], 'Do not change the some_numbers list!'
assert len(positive) == 4, 'The length of the positive list is not correct.'
assert sum(positive) == 58, 'The elements of the positive list are not correct.'
assert sum(other) == -111, 'The elements of the other list are not correct.'
print(positive)
print(other)
'Well done!'
```

### Exercise 2.4
Below is a for loop that assigns the elements of the list `some_more_numbers` to the lists `my_numbers` and `your_numbers` one by one. Both lists are printed at the end. Run the code if you like to see what it does.

Add an `if-else` statement inside the `for` loop which filters the elements of the `some_more_numbers` list so that numbers divisible by 3 and larger than 20 are assigned to the `my_numbers` list and all other numbers are assigned to the `your_numbers` list.

**Do not remove the code that is already in the cell!** Just add the `if-else` statement where indicated and adjust the indenting if necessary.

```python
# Rerun this cell if you lose the some_more_numbers list
some_more_numbers = [-3, 9, 21, 172, 200, -51, 81, -12, 3000, 18, 702]
```

<!-- grade_id: cell-14ccc85468cfe710, points: 0, grade: false, solution: true -->
```python
some_more_numbers = [-3, 9, 21, 172, 200, -51, 81, -12, 3000, 18, 702]
my_numbers = []
your_numbers = []

for number in some_more_numbers:
    # add the if-else statement here
    my_numbers.append(number)
    your_numbers.append(number)

print(my_numbers)
print(your_numbers)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-49001e749ed9ddc7, points: 1, grade: true, solution: false -->
```python
assert some_more_numbers == [-3, 9, 21, 172, 200, -51, 81, -12, 3000, 18, 702], 'Do not change the some_numbers list!'
assert len(my_numbers) == 4, 'The length of the my_numbers list is not correct.'
assert len(your_numbers) == 7, 'The length of the your_numbers list is not correct.'
assert sum(my_numbers) == 3804, 'The elements of the my_numbers list are not correct.'
assert sum(your_numbers) == 333, 'The elements of the your_numbers list are not correct.'
print(my_numbers)
print(your_numbers)
'Well done!'
```

### Exercise 2.5
Below is a for loop that assigns the elements of the list `even_more_numbers` to the lists `list_7_22` and `list_other` one by one. Both lists are printed at the end. Run the code if you like to see what it does.

Add an `if-else` statement inside the `for` loop which filters the elements of the `even_more_numbers` list so that numbers divisible by 7 or smaller than 22 are assigned to the `list_7_22` list and all other numbers are assigned to the `list_other` list.

**Do not remove the code that is already in the cell!** Just add the `if-else` statement where indicated and adjust the indenting if necessary.

```python
# Rerun this cell if you lose the even_more_numbers list
even_more_numbers = [22, 14, -3, -81, 55, 84, 123, 75, 1024, 308, 13]
```

<!-- grade_id: cell-18d279a61e25d285, points: 0, grade: false, solution: true -->
```python
even_more_numbers = [22, 14, -3, -81, 55, 84, 123, 75, 1024, 308, 13]
list_7_22 = []
list_other = []

for number in even_more_numbers:
    # add the if-else statement here
    list_7_22.append(number)
    list_other.append(number)

print(list_7_22)
print(list_other)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-307eb9e5088602d0, points: 1, grade: true, solution: false -->
```python
assert even_more_numbers == [22, 14, -3, -81, 55, 84, 123, 75, 1024, 308, 13], 'Do not change the some_numbers list!'
assert len(list_7_22) == 6, 'The length of the list_7_22 list is not correct.'
assert len(list_other) == 5, 'The length of the list_other list is not correct.'
assert sum(list_7_22) == 335, 'The elements of the list_7_22 list are not correct.'
assert sum(list_other) == 1299, 'The elements of the list_other list are not correct.'
print(list_7_22)
print(list_other)
'Well done!'
```

## 3. For loops

### Exercise 3.1
Below is an empty list called `numbers_0_99`. Write a `for` loop which will fill the list with integers from 0 to 99. The filled list should look like this `[0, 1, 2, ..., 89, 99]`.

<!-- grade_id: cell-279331d5b7558fc6, points: 0, grade: false, solution: true -->
```python
numbers_0_99 = []
# for ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-67793628a54caefd, points: 1, grade: true, solution: false -->
```python
assert len(numbers_0_99) == 100, 'The length of the numbers_0_99 list is not correct.'
assert sum(numbers_0_99) == 4950, 'The elements of the numbers_0_99 list are not correct.'
assert numbers_0_99[0] == 0, 'The elements of the numbers_0_99 list are not correct.'
assert numbers_0_99[52] == 52, 'The elements of the numbers_0_99 list are not correct.'
assert numbers_0_99[99] == 99, 'The elements of the numbers_0_99 list are not correct.'
'Well done!'
```

### Exercise 3.2
Below is an empty list called `numbers_12_123`. Write a `for` loop which will fill the list with integers from 12 to 123. The filled list should look like this `[12, 13, 14, ..., 122, 123]`.

<!-- grade_id: cell-51e6fc266a201d76, points: 0, grade: false, solution: true -->
```python
numbers_12_123 = []
#for ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5f545c9ef1951a13, points: 1, grade: true, solution: false -->
```python
assert len(numbers_12_123) == 112, 'The length of the numbers_12_123 list is not correct.'
assert sum(numbers_12_123) == 7560, 'The elements of the numbers_12_123 list are not correct.'
assert numbers_12_123[0] == 12, 'The elements of the numbers_12_123 list are not correct.'
assert numbers_12_123[52] == 64, 'The elements of the numbers_12_123 list are not correct.'
assert numbers_12_123[111] == 123, 'The elements of the numbers_12_123 list are not correct.'
'Well done!'
```

### Exercise 3.3
Below is an empty list called `numbers_by_3`. Write a `for` loop which will fill the list with integers from 4 to 118, but only includes every third number. The filled list should look like this `[4, 7, 10, ..., 115, 118]`.

<!-- grade_id: cell-9e995e2964022bb8, points: 0, grade: false, solution: true -->
```python
numbers_by_3 = []
#for ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0948ff2530ac8ff1, points: 1, grade: true, solution: false -->
```python
assert len(numbers_by_3) == 39, 'The length of the numbers_by_3 list is not correct.'
assert sum(numbers_by_3) == 2379, 'The elements of the numbers_by_3 list are not correct.'
assert numbers_by_3[0] == 4, 'The elements of the numbers_by_3 list are not correct.'
assert numbers_by_3[1] == 7, 'The elements of the numbers_by_3 list are not correct.'
assert numbers_by_3[37] == 115, 'The elements of the numbers_by_3 list are not correct.'
assert numbers_by_3[38] == 118, 'The elements of the numbers_by_3 list are not correct.'
'Well done!'
```

### Exercise 3.4
Below is an empty list called `backwards_by_7`. Write a `for` loop which will fill the list with integers from 1024 to 513, but only includes every seventh number. The filled list should look like this `[1024, 1017, 1010, ..., 520, 513]`.

<!-- grade_id: cell-6fbebfe5181acd88, points: 0, grade: false, solution: true -->
```python
backwards_by_7 = []
#for ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-039aa2c4b8032dd6, points: 1, grade: true, solution: false -->
```python
assert len(backwards_by_7) == 74, 'The length of the backwards_by_7 list is not correct.'
assert sum(backwards_by_7) == 56869, 'The elements of the backwards_by_7 list are not correct.'
assert backwards_by_7[0] == 1024, 'The elements of the backwards_by_7 list are not correct.'
assert backwards_by_7[1] == 1017, 'The elements of the backwards_by_7 list are not correct.'
assert backwards_by_7[72] == 520, 'The elements of the backwards_by_7 list are not correct.'
assert backwards_by_7[73] == 513, 'The elements of the backwards_by_7 list are not correct.'
'Well done!'
```

## Exercise 4: While loops
You're gonna hate me now, but trust me, I know what's good for you. :) We will repeat the same exercises as in exercise 3, but with `while` loops. Remember to initialize and increment your control variables.

<img src="./media/danger.png" width = "100"/>

<left> High potential for infinite loops!</left>

### Exercise 4.1
Below is an empty list called `while_0_99`. Write a `while` loop which will fill the list with integers from 0 to 99. The filled list should look like this `[0, 1, 2, ..., 89, 99]`. 

<!-- grade_id: cell-a5e33f52f7b7bbdc, points: 0, grade: false, solution: true -->
```python
while_0_99 = []
# while ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c2ba61f4bfb8be0a, points: 1, grade: true, solution: false -->
```python
assert len(while_0_99) == 100, 'The length of the while_0_99 list is not correct.'
assert sum(while_0_99) == 4950, 'The elements of the while_0_99 list are not correct.'
assert while_0_99[0] == 0, 'The elements of the while_0_99 list are not correct.'
assert while_0_99[52] == 52, 'The elements of the while_0_99 list are not correct.'
assert while_0_99[99] == 99, 'The elements of the while_0_99 list are not correct.'
'Well done!'
```

### Exercise 4.2
Below is an empty list called `while_12_123`. Write a `while` loop which will fill the list with integers from 12 to 123. The filled list should look like this `[12, 13, 14, ..., 122, 123]`.

<!-- grade_id: cell-cd6c57af87d60a87, points: 0, grade: false, solution: true -->
```python
while_12_123 = []
#while ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3026effaf79903db, points: 1, grade: true, solution: false -->
```python
assert len(while_12_123) == 112, 'The length of the while_12_123 list is not correct.'
assert sum(while_12_123) == 7560, 'The elements of the while_12_123 list are not correct.'
assert while_12_123[0] == 12, 'The elements of the while_12_123 list are not correct.'
assert while_12_123[52] == 64, 'The elements of the while_12_123 list are not correct.'
assert while_12_123[111] == 123, 'The elements of the while_12_123 list are not correct.'
'Well done!'
```

### Exercise 4.3
Below is an empty list called `while_by_3`. Write a `while` loop which will fill the list with integers from 4 to 118, but only includes every third number. The filled list should look like this `[4, 7, 10, ..., 115, 118]`.

<!-- grade_id: cell-355d5786597dc995, points: 0, grade: false, solution: true -->
```python
while_by_3 = []
#while ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-74c07725c5a92300, points: 1, grade: true, solution: false -->
```python
assert len(while_by_3) == 39, 'The length of the while_by_3 list is not correct.'
assert sum(while_by_3) == 2379, 'The elements of the while_by_3 list are not correct.'
assert while_by_3[0] == 4, 'The elements of the while_by_3 list are not correct.'
assert while_by_3[1] == 7, 'The elements of the while_by_3 list are not correct.'
assert while_by_3[37] == 115, 'The elements of the while_by_3 list are not correct.'
assert while_by_3[38] == 118, 'The elements of the while_by_3 list are not correct.'
'Well done!'
```

### Exercise 4.4
Below is an empty list called `while_backwards_by_7`. Write a `while` loop which will fill the list with integers from 1024 to 513, but only includes every seventh number. The filled list should look like this `[1024, 1017, 1010, ..., 520, 513]`.

<!-- grade_id: cell-3f422105e028174b, points: 0, grade: false, solution: true -->
```python
while_backwards_by_7 = []
#while ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-83cc31713646e521, points: 1, grade: true, solution: false -->
```python
assert len(while_backwards_by_7) == 74, 'The length of the while_backwards_by_7 list is not correct.'
assert sum(while_backwards_by_7) == 56869, 'The elements of the while_backwards_by_7 list are not correct.'
assert while_backwards_by_7[0] == 1024, 'The elements of the while_backwards_by_7 list are not correct.'
assert while_backwards_by_7[1] == 1017, 'The elements of the while_backwards_by_7 list are not correct.'
assert while_backwards_by_7[72] == 520, 'The elements of the while_backwards_by_7 list are not correct.'
assert while_backwards_by_7[73] == 513, 'The elements of the while_backwards_by_7 list are not correct.'
'Well done!'
```

## Exercise 5: More loops
You can use `for` or `while` loops in these exercises, whatever works better.

### Exercise 5.1 Broken words
We have two lists with pieces of words: `start_list` and `end_list`. You want to put the words back together. Add each word part from the `start_list` with each word part from the `end_list` and assign the results to the list `whole_list`.

Example: the first word in the `start_list` is 'beauty' and in the `end_list` it's 'ful', so the first word in the whole_list will be 'beauty' + 'ful' which is 'beautiful'.

You will need more than one loop here.

<!-- grade_id: cell-b1fe5b37cf75dd7d, points: 0, grade: false, solution: true -->
```python
start_list = ['beauty', 'cat', 'ice']
end_list = ['ful', 'less', 'thon', 'house', 'woman', 'tea']
#whole_list = []

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-dda0109e4ca4c20f, points: 1, grade: true, solution: false -->
```python
assert len(whole_list) == 18, 'The length of whole_list is not correct.'
assert hashlib.sha1(bytes(''.join(sorted(whole_list)), encoding='utf8')).hexdigest() == \
'883793d9abeef55b381707a329b4f0ccbaf1e5ca', 'The elements of whole_list are not correct.'
'Well done!'
```

### Exercise 5.2
You have a dictionary of things called `household`. The keys of the dictionary are the things and the values indicate if the thing is new or old. Loop over the things in the dictionary and sort them into two lists, `new_things` and `old_things` according to their state.

Example: 'fridge' is 'old', so you append 'fridge' to the `old_things` list.

<!-- grade_id: cell-9a00daaa64a76117, points: 0, grade: false, solution: true -->
```python
household = {'fridge':'old', 'oven':'new', 'bed':'new', 'table':'old',
             'bookcase':'new', 'stove':'old', 'piano':'old'}
new_things = []
old_things = []

# YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: cell-a396380e74e0c17b, points: 1, grade: true, solution: false -->
```python
assert len(old_things) == 4, 'The length of the old_things list is not correct.'
assert hashlib.sha1(bytes(''.join(sorted(old_things)), encoding='utf8')).hexdigest() == \
'73018c4decff74d0bd40a3d99f94abcba64a7d60', 'The elements of the old_things list are not correct.'
assert len(new_things) == 3, 'The length of the new_things list is not correct.'
assert hashlib.sha1(bytes(''.join(sorted(new_things)), encoding='utf8')).hexdigest() == \
'7014c84e5f7934a4e50152c695262e7a143ca434', 'The elements of the new_things list are not correct.'
'Well done!'
```

### Exercise 5.3 Game of cards
We'll be playing a game of cards. You have a function called `card_deck` that will return a card with a number value between 0 and 14. The game rules are as follows:
In every round, you take a card from the deck:

```python
card = card_deck()`
```

You add the card to your `hand` list:

```python
hand.append(card)
```

If the sum of the values of the cards in the `hand` list exceeds 51, the game ends. Starting from an empty hand, write a loop that simulates this game.

<!-- grade_id: cell-b4476e7929e1d76a, points: 0, grade: false, solution: true -->
```python
hand = []
#card = card_deck()
#while ...

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-2f1bea06088caf5e, points: 1, grade: true, solution: false -->
```python
assert sum([n>=0 and n<= 14 and isinstance(n,int) for n in hand]) == len(hand),\
'There are cards with invalid values in your hand.'
assert sum(hand) > 51, 'The sum of cards in your hand is not larger than 51.'
assert sum(hand[:-1]) <= 51, 'The sum of all but last card in your hand is larger than 51'
```

Congratulations, you made it! Now that you can write a for loop with your eyes closed, you can move on to more interesting exercises in the Exercise notebook 1. They are also a bit more difficult, so remember to plan your workflow before writing the code.

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2024/blob/main/weekly-workflow.md#link-to-grading)