# SLU05 - Exercise notebook 2

Welcome to the world of `Python Flow Control`! 

Developing a thorough grasp of the concepts and techniques in this topic will provide you with a good programming foundation. We have included many exercises here, some are challenging. Take your time and practice. We are happy to assist, and if you don't have time to complete everything within this week, you can always return to it later.
At any time when you are stuck feel free to use the function "hints" that is already imported for you in the beginning. Call help(hints) if you would like to know how it works!

## Start by importing these packages

```python
# Just for evaluating the results.
import math
import json
import hashlib
from utils import hints
```

```python
help (hints)
```

Output:
```
Help on function hints in module utils:

hints(exercise, hint_number)
    hints (exercise, hint_number)
    exercise: string with exercise number/chapter for example "1.1"
    hint_number: int, number of hint,  for example 1  - there are up to 3 for each exercise
```

## Part 1 - Conditionals & Boolean Algebra

Welcome to Planet X! 

Every 7 years the playmakers get together to create Game X.  Candidates from around the planet prepare to entertain the world. The candidates and their characteristics are shown below:


| Candidates | Age | Height (cm) | Weight (kg) | Strength |
| - |-: |-: |-: |-: |
|Squindle|213|1810|954|6|
|Flort|73|1225|769|8|
|Hambula|82|999|642|4|
|Kistapho|942|1829|842|6|
|Vilawi|636|819|531|5|
|Anhaular|231|1681|3762|9|
|Qintari|60|1732|522|4|
|Wendu|7|3|5|2|
|Baxla|65|875|464|4|
|Uluetto|251|1362|3647|2|
|Reomult|740|975|539|9|


### 1.1) Create Candidates Dictionary

Create a nested dictionary of the candidates, `candidate_dict`, using the candidate´s name as the top-level key to capturing all the above characteristics. The second level keys should be `"age"`, `"height"`, `"weight"` and `"strength"`. 

For instance, accessing `candidate_dict["Squindle"]["age"]` should return `213`.

Please double check that the entries are correct as this dictionary will be used for the remaining exercises in this notebook.

Hint: below is an example of a nested dictionary. It is a dictionary of dictionaries.
    
    nested_dict = { 'dict1': {'key_A': 'value_A'},
                    'dict2': {'key_B': 'value_B'}}

<!-- grade_id: cell-458be56634b57ca9, points: 0, grade: false, solution: true -->
```python
# Create a dictionary with the candidates characteristics above
# Assign the dictionary to variable candidate_dict.

# candidate_dict = ...

### BEGIN SOLUTION

# Here is the actual solution
candidate_dict = {
        'Squindle': {'age': 213, 'height': 1810, 'weight': 954, 'strength': 6},
        'Flort': {'age': 73, 'height': 1225, 'weight': 769, 'strength': 8},
        'Hambula': {'age': 82, 'height': 999, 'weight': 642, 'strength': 4},
        'Kistapho': {'age': 942, 'height': 1829, 'weight': 842, 'strength': 6},
        'Vilawi': {'age': 636, 'height': 819, 'weight': 531, 'strength': 5},
        'Anhaular': {'age': 231, 'height': 1681, 'weight': 3762, 'strength': 9},
        'Qintari': {'age': 60, 'height': 1732, 'weight': 522, 'strength': 4},
        'Wendu': {'age': 7, 'height': 3, 'weight': 5, 'strength': 2},
        'Baxla': {'age': 65, 'height': 875, 'weight': 464, 'strength': 4},
        'Uluetto': {'age': 251, 'height': 1362, 'weight': 3647, 'strength': 2},
        'Reomult': {'age': 740, 'height': 975, 'weight': 539, 'strength': 9},
}

### END SOLUTION

```

<!-- grade_id: cell-68331a27a8208436, points: 2, grade: true, solution: false -->
```python
assert isinstance(candidate_dict, dict), "Review the data type of candidate_dict."
assert len(candidate_dict) == 11, "The size of the dictionary is incorrect."
assert set(candidate_dict.keys()) == \
  set(['Squindle', 'Flort', 'Hambula', 'Kistapho', 'Vilawi', 'Anhaular', 'Qintari', 'Wendu', 'Baxla', 'Uluetto', 'Reomult']),\
  "The dictionary has incorrect candidate names."

assert sum([len(candidate_dict[k]) for k in candidate_dict.keys()]) == 44, 'Some of the nested candidate dictionaries have incorrect length.'
for candidate in candidate_dict.keys():
    assert set(candidate_dict[candidate].keys()) == set(['age', 'height', 'strength', 'weight']), \
    f"The {candidate} dictionary has mismatched candidate characteristics."

assert len([isinstance(candidate_dict[k][c],int) for k in candidate_dict.keys() for c in ['age', 'height', 'strength', 'weight']]) == 44, \
'The values of age, height, strength, weight should be integers.'

candidates = ['Squindle', 'Flort', 'Hambula', 'Kistapho', 'Vilawi', 'Anhaular', 'Qintari', 'Wendu', 'Baxla', 'Uluetto', 'Reomult']
hash_list = ['f1a2dd2904a3a3051f734189e008fd59bb151cbb3c88625ed6da12548d2e2d34', '0fa60da40f6431ad2aad3f039f52a5e51af3228c5ef2e5ca9e6648077fd65f7e',
 '2cbdd1689217081ebb514ad1de632617f0686e73d2d5a9f1f39bf86fe3413fcd', '1616836491c36ac4a55fde630f414fbd512e10daccd2397e74ddae761e37f687',
 'bc89e34e334866ac63cbffbe7179805493529cdace179e87ff094bc9784439cd', 'd9b5d379089aec1946b613f5991ca2338772a4fc10833aafa435d68d22ec8dce',
 '41ea639e8a5ea8261f705b011dc5bb67d2c20f941131efc62174b68e5180d9b4', '879ae9a0a821ab4f1c99b2d97097a6405844ebc5cb72592b6ac17da5995f652c',
 '7085339eda3f735877f4e19551d6f5c3e112efa9eb09cad6bcb73987aa332cf0', '3e3c8d03af0958c541b986b5d10d659e007924b608f43fa444108318f7e0e0a0',
 'd89c91c710824e013815377b23a48b038b7d2231cd3423795479bd2966dededf']
hash_dict=dict(zip(candidates,hash_list))
for candidate in candidate_dict.keys():
    assert hashlib.sha256(
        json.dumps(''.join([str(candidate_dict[candidate][k]) for k in ['age', 'height', 'strength', 'weight']])).encode()).hexdigest() ==\
       hash_dict[candidate], f'The values in the {candidate} dictionary are not correct.'
print("Your solution appears correct!")
```

Output:
```
Your solution appears correct!
```

### 1.2)  Order of precedence

To compete in Game X, one must be able to program in Python! 

In the pre-screening round, candidates are required to arrange a set of stones (A to G) in the correct sequence in order to unlock a door that holds their entrance ticket to the Game. 

Candidates are asked to place the stones in a sequence following the order of **operation evaluation** in the following boolean expression in Python. If in doubt, refer to the `table of precedence` to identify the order in which the operations are executed. Each operator is identified with a letter from `A` to `G`.

The solution should be a **list** where the first element is the letter of the first operation to be performed, the second element is the second operation to be performed, and so on until the last operation to be performed.

The result should resemble something like `["A", "B", "D", "G", ...]`

The expression is:

```python
#You don't need to execute this cell!

variable = 3 != 1+1 and (2**3 > 1 or False)
#        A   B   C   D    E   F   G
```

The letters that identify the operations are:
- A: `=`
- B: `!=`
- C: `+`
- D: `and`
- E: `**`
- F: `>`
- G: `or`

<!-- grade_id: cell-2910243951622497, points: 0, grade: false, solution: true -->
```python
# Create a list named operators_list that contains the identifiers of the operators sorted by the execution order.
# Make sure that you type the strings exactly.

# operators_list = ...
### BEGIN SOLUTION
# Here is the actual solution
operators_list = ["C", "B", "E", "F", "G", "D", "A"]

### END SOLUTION
```

<!-- grade_id: cell-4a271c9901ebcdd3, points: 1, grade: true, solution: false -->
```python
assert isinstance(operators_list, list), "Review the data type of operators_list."
assert len(operators_list) == 7, "The number of strings is incorrect."
assert hashlib.sha256(json.dumps(''.join(operators_list).upper()).encode()).hexdigest() == \
'1ffc8471817a2bf2b5d0b24e0b799a943dd5d88be520ed09b052cdbe70e487b3', 'The operator order is not correct.'
print("Your solution is correct!")
```

Output:
```
Your solution is correct!
```

### 1.3) if statements¶

The candidates' colors are assigned using the following Python code. 

<img src="./media/if_code.png" />

Please use your logic, rather than executing the code, and decide which colors are assigned to the following players:

    *	Kistapho
    *	Vilawi
    *	Baxla
    *	Qintari

<!-- grade_id: cell-86b889f5f10ed961, points: 0, grade: false, solution: true -->
```python
# Assign the color for each candidate resulting from the code above to these variables:

# kistapho_color = ...
# vilawi_color = ...
# baxla_color = ...
# qintari_color = ...

### BEGIN SOLUTION

# Here is the actual solution
kistapho_color = "blue"
vilawi_color = "purple"
baxla_color = "yellow"
qintari_color = "red"

### END SOLUTION
```

<!-- grade_id: cell-a68470c659fe1a28, points: 1, grade: true, solution: false -->
```python
assert isinstance(kistapho_color, str), "Variable kistapho_color should be a string."
assert isinstance(vilawi_color, str), "Variable vilawi_color should be a string."
assert isinstance(baxla_color, str), "Variable baxla_color should be a string."
assert isinstance(qintari_color, str), "Variable qintari_color should be a string."

assert hashlib.sha256(json.dumps(kistapho_color).encode()).hexdigest() == '3aae004f448787fcf06c2d53933fff00a45fc5117c21e1070f967e37d43760ef', "The value of kistapho_color is incorrect."
assert hashlib.sha256(json.dumps(vilawi_color).encode()).hexdigest() == 'e3c587c5f6c356d0d488392bc278fe2c1732f8029f311508e48df180e5ed8356', "The value of vilawi_color is incorrect."
assert hashlib.sha256(json.dumps(baxla_color).encode()).hexdigest() == 'c72b530ab18e2b7c06129203771c7e5ae28618a981c971046e47fb19693a6f23', "The value of baxla_color is incorrect."
assert hashlib.sha256(json.dumps(qintari_color).encode()).hexdigest() == '4bc9274d8758b867e8c071a6082c68a2eb1668fc5ad34f631f96bf94f354c1e1', "The value of qintari_color is incorrect."
print("Your solution is correct!")
```

Output:
```
Your solution is correct!
```

## Part 2 - While & For Loops

In order to qualify for Game X, the candidates must meet the following requirements:
- Height >= 800 cm
- 5 < BMI < 20


### 2.1) Calculate BMI

Read candidates' characteristics from `candidate_dict`, and calculate their `BMI` using the following formula. 

Update the dictionary with their BMI using the key `bmi`. **Pay attention to the unit**, convert if necessary.
    
$$BMI = \frac{Weight}{Height^{2}} $$
             
                                                                where
                                                                    Weight is measured in kg
                                                                    Height is measured in meters


<!-- grade_id: cell-e1b6130c39f10570, points: 0, grade: false, solution: true -->
```python
# Calculate BMI for each candidate and add the results to candidate_dict.

### BEGIN SOLUTION

# Here is the actual solution
for person in candidate_dict.values():
    person["bmi"] = person["weight"] / ((person["height"] / 100) **2)
    
### END SOLUTION    
```

<!-- grade_id: cell-be5348ea4d11a7d7, points: 2, grade: true, solution: false -->
```python
assert sum([len(candidate_dict[k]) for k in candidate_dict.keys()]) == 55, 'Some of the nested candidate dictionaries have incorrect length.'
for candidate in candidate_dict.keys():
    assert set(candidate_dict[candidate].keys()) == set(['age', 'height', 'strength', 'weight', 'bmi']), \
    f"The {candidate} dictionary has mismatched candidate characteristics."

assert math.isclose(candidate_dict["Squindle"]["bmi"], 2.91, abs_tol=0.01), "The BMI for Squindle is incorrect."
assert math.isclose(candidate_dict["Flort"]["bmi"], 5.12, abs_tol=0.01), "The BMI for Flort is incorrect."
assert math.isclose(candidate_dict["Hambula"]["bmi"], 6.43, abs_tol=0.01), "The BMI for Hambula is incorrect."
assert math.isclose(candidate_dict["Kistapho"]["bmi"], 2.52, abs_tol=0.01), "The BMI for Kistapho is incorrect."
assert math.isclose(candidate_dict["Vilawi"]["bmi"], 7.92, abs_tol=0.01), "The BMI for Vilawi is incorrect."
assert math.isclose(candidate_dict["Anhaular"]["bmi"], 13.31, abs_tol=0.01), "The BMI for Anhaular is incorrect."
assert math.isclose(candidate_dict["Qintari"]["bmi"], 1.74, abs_tol=0.01), "The BMI for Qintari is incorrect."
assert math.isclose(candidate_dict["Wendu"]["bmi"], 5555.56, abs_tol=0.01), "The BMI for Wendu is incorrect."
assert math.isclose(candidate_dict["Baxla"]["bmi"], 6.06, abs_tol=0.01), "The BMI for Baxla is incorrect."
assert math.isclose(candidate_dict["Uluetto"]["bmi"], 19.66, abs_tol=0.01), "The BMI for Uluetto is incorrect."
assert math.isclose(candidate_dict["Reomult"]["bmi"], 5.67, abs_tol=0.01), "The BMI for Reomult is incorrect."

print("Your solution is correct!")
```

Output:
```
Your solution is correct!
```

### 2.2) Power Potion to Increase Height

Baby `Wendu` really wants to compete in the game. She is very advanced for her age, but her height doesn´t meet the minimum requirement for Game X (i.e. `height >= 800cm`). Luckily there is a power potion that she can inhale that will double her height every minute. 

Use a loop to calculate the `minimum time` she needs to inhale the power potion to reach the height requirement. Update the `candidates dictionary` with her new `height` and `BMI`. 

Note - If you want to re-execute the cell after it completes successfully the first time around, you may need to reset the dictionary by re-running the dictionary creation cell in exercise 1.1

<!-- grade_id: cell-64700ac1d2373664, points: 0, grade: false, solution: true -->
```python
# Calculate the minimum time for Wendu to reach the required height for Game X, 
# and store it in the variable time_taken
# update the candidate dictionary for Wendu

# time_taken = ...
# candidate_dict["Wendu"]["height"] = ...
# candidate_dict["Wendu"]["bmi"] = ...

### BEGIN SOLUTION

# Here is the actual solution

wendu_height = candidate_dict["Wendu"]["height"]
time_taken = 0

while wendu_height < 800:
    wendu_height = wendu_height * 2
    time_taken += 1

wendu_bmi = candidate_dict["Wendu"]["weight"] / ((wendu_height / 100) **2)

candidate_dict["Wendu"]["height"] = wendu_height
candidate_dict["Wendu"]["bmi"] = wendu_bmi

### END SOLUTION
```

<!-- grade_id: cell-96095f682edc0492, points: 2, grade: true, solution: false -->
```python
assert isinstance(candidate_dict["Wendu"]["height"], int), "Review the data type of Wendu's height."
assert isinstance(candidate_dict["Wendu"]["bmi"], float), "Review the data type of Wendu's BMI."
assert isinstance(time_taken, int), "Review the data type of the time taken."

assert hashlib.sha256(json.dumps(time_taken).encode()).hexdigest() == '19581e27de7ced00ff1ce50b2047e7a567c76b1cbaebabe5ef03f7c3017bb5b7', "The value of time taken is incorrect."
assert hashlib.sha256(json.dumps(candidate_dict["Wendu"]["height"]).encode()).hexdigest() == 'b51e45a12fbae3d0ee2bf77f1a4f80cbf642e2b4d1c237d2c0f7053a54f6b388', "The value of Wendu's height is incorrect."
assert math.isclose(candidate_dict["Wendu"]["bmi"], 0.021, abs_tol=0.001), "The BMI for Wendu is incorrect."
print(f'Wendu\'s new values: height {candidate_dict["Wendu"]["height"]}, bmi {candidate_dict["Wendu"]["bmi"]}')
print("Your solution is correct!")
```

Output:
```
Wendu's new values: height 1536, bmi 0.021192762586805556
Your solution is correct!
```

### 2.3) Balls of Hopstow to Increase Weight

Now baby Wendu needs to gain some weight, as her BMI falls well below the requirement (`5 < BMI < 20`). 

In Planet X, eating `Balls of Hopstow` increases a person’s weight by 10 kg. How many balls of Hopstow does Wendu need to eat, in order to enter the game? Update the `candidates dictionary` with her new `weight` and her new `BMI`.

```python
# This cell contains the original characteristics of Wendu.
# Do not run it after executing your solution,
# unless you want to reset to the original values.

wendu_weight_original =  candidate_dict["Wendu"]["weight"]
wendu_bmi_original =  candidate_dict["Wendu"]["bmi"]

```

<!-- grade_id: cell-cb72287d39f4d7da, points: 0, grade: false, solution: true -->
```python
# Calculate the number of balls of Hopstow Wendu needs to consume, and store it in the variable num_balls.
# keep track of Wendu's BMI changes by storing each value in bmi_list, starting with her original BMI
# update candidate_dict with Wendu's new values

# candidate_dict["Wendu"]["weight"] = wendu_weight_original
# candidate_dict["Wendu"]["bmi"] = wendu_bmi_original

# bmi_list = [...]
# num_balls = ...

### BEGIN SOLUTION

# Here is the actual solution
wendu_weight =  wendu_weight_original
wendu_bmi =  wendu_bmi_original

required_bmi = 5
num_balls = 0
bmi_list = [wendu_bmi]

while wendu_bmi < required_bmi:
    wendu_weight += 10
    wendu_bmi = wendu_weight / ((wendu_height / 100)**2)
    bmi_list.append(wendu_bmi)
    num_balls += 1
    
candidate_dict["Wendu"]["weight"] = wendu_weight
candidate_dict["Wendu"]["bmi"] = wendu_bmi

### END SOLUTION
```

<!-- grade_id: cell-329ed1829c380112, points: 2, grade: true, solution: false -->
```python
assert isinstance(num_balls, int), "Review the data type of num_balls."
assert isinstance(candidate_dict["Wendu"]["weight"], int), "Review the data type of Wendu's weight."
assert isinstance(candidate_dict["Wendu"]["bmi"], float), "Review the data type of Wendu's BMI."

assert hashlib.sha256(json.dumps(num_balls).encode()).hexdigest() == '85daaf6f7055cd5736287faed9603d712920092c4f8fd0097ec3b650bf27530e', "The value of num_balls is incorrect."
assert hashlib.sha256(json.dumps(len(bmi_list)).encode()).hexdigest() == "3038bfb575bee6a0e61945eff8784835bb2c720634e42734678c083994b7f018", "The number of items in bmi_list is incorrect."
assert hashlib.sha256(json.dumps(candidate_dict["Wendu"]["weight"]).encode()).hexdigest() == "efec9aaf21433bf806e7681de337cac7dbecfbf17b22ad3bcdfe5e46e564f32f", "The weight for Wendu is incorrect."

assert math.isclose(candidate_dict["Wendu"]["bmi"], 5.02, abs_tol=0.01), "The BMI for Wendu is incorrect."
print(f'Wendu\'s new values: weight {candidate_dict["Wendu"]["weight"]}, bmi {candidate_dict["Wendu"]["bmi"]}')
print(f'bmi_list {bmi_list}')
print(f'number of balls eaten {num_balls}')
print("Your solution is correct!")
```

Output:
```
Wendu's new values: weight 1185, bmi 5.022684733072917
bmi_list [0.021192762586805556, 0.06357828776041667, 0.10596381293402778, 0.1483493381076389, 0.19073486328125, 0.2331203884548611, 0.2755059136284722, 0.3178914388020833, 0.3602769639756945, 0.4026624891493056, 0.4450480143229167, 0.4874335394965278, 0.529819064670139, 0.57220458984375, 0.6145901150173612, 0.6569756401909722, 0.6993611653645834, 0.7417466905381944, 0.7841322157118056, 0.8265177408854167, 0.8689032660590278, 0.911288791232639, 0.95367431640625, 0.9960598415798612, 1.0384453667534723, 1.0808308919270833, 1.1232164171006944, 1.1656019422743056, 1.2079874674479167, 1.250372992621528, 1.2927585177951388, 1.33514404296875, 1.3775295681423612, 1.4199150933159723, 1.4623006184895835, 1.5046861436631944, 1.5470716688368056, 1.5894571940104167, 1.631842719184028, 1.6742282443576388, 1.71661376953125, 1.7589992947048612, 1.8013848198784723, 1.8437703450520835, 1.8861558702256944, 1.9285413953993056, 1.9709269205729167, 2.013
```

### 2.4) Calculate the Prize Pool

The prize pool for Game X consists of `100 coins`, where the value of each coin `increases by 3` incrementally, i.e. the first coin is worth 1 point, the second is worth 4 points, the third is worth 7 points, etc. 

What is the total value of the prize pool? 

<!-- grade_id: cell-3243cc018c19ef15, points: 0, grade: false, solution: true -->
```python
# store the value in the variable prize_pool

# prize_pool = ...

### BEGIN SOLUTION

# Here is the actual solution
prize_pool = 0
next_val = 1

for i in range(100):
    prize_pool += next_val
    next_val += 3
    
### END SOLUTION
```

<!-- grade_id: cell-40809f394cee6d36, points: 2, grade: true, solution: false -->
```python
assert isinstance(prize_pool, int), "Review the data type of the prize_pool."
assert hashlib.sha256(json.dumps(prize_pool).encode()).hexdigest() == '4bd6fb30bbbaba340fb0267845a29e13ac60300746152fb9ce3fc7434586e207', "The value of prize_pool is incorrect."
print(f'prize pool {prize_pool}')
print("Your solution is correct!")

```

Output:
```
prize pool 14950
Your solution is correct!
```

## Part 3 - Using Conditionals and Loops together 

### 3.1) Identify the Players that Qualify

Using the `candidates dictionary` above, which of the candidates now satisfy the following criteria and thus qualify for Game X? 
*    Height >= 800 cm
*    5 < BMI < 20

Use a loop and store their names in the `player_list`, and then sort them in `alphabetical order`.

<!-- grade_id: cell-07713da08525370c, points: 0, grade: false, solution: true -->
```python
# Create a list named player_list to store the names of players that qualify for Game X

# player_list = [...]

### BEGIN SOLUTION

player_list = []

for name, attribute in candidate_dict.items():
    if (attribute["height"] >= 800) and (attribute["bmi"] > 5) and (attribute["bmi"] < 20):
        player_list.append(name)

player_list = sorted(player_list)

### END SOLUTION
```

<!-- grade_id: cell-9477efc58181088a, points: 2, grade: true, solution: false -->
```python
assert isinstance(player_list, list), "Variable player_list should be a list."
assert hashlib.sha256(json.dumps(len(player_list)).encode()).hexdigest() == '2c624232cdd221771294dfbb310aca000a0df6ac8b66b696d90ef06fdefb64a3', "The number of items in player_list is incorrect."
assert hashlib.sha256(json.dumps(''.join(player_list)).encode()).hexdigest() == 'cfaca8aace4a00f2274aabc8062b1c8b0515e5991d0af06560889d7afbe46778', 'At least one player is incorrect or the list in not sorted well.'
print(f'player_list {player_list}')
print("Your solution is correct!")
```

Output:
```
player_list ['Anhaular', 'Baxla', 'Flort', 'Hambula', 'Reomult', 'Uluetto', 'Vilawi', 'Wendu']
Your solution is correct!
```

### 3.2) Climb that Wall

Now that we have a `list of the players`, the game begins! The first challenge is a climbing competition. The players will be given `5 minutes` to ascend the `Wall to Nowhere`. 

Unfortunately their climbing rate is proportional to their height. Players shorter than `1200 cm` can climb **$\frac{1}{2}$ of their height** per minute while the other player climbs at the rate of **$\frac{1}{3}$ of their height** per minute. 

As a show of respect for the elders, players older than `500 years` start at `15 meters` above ground.

This can be solved using a simple mathematical equation, however, you can also try to implement it as a `for loop` challenge.

Store the result of each player in the dictionary `climb_dict_5min`, with the player names being the top level key. ie. `climb_dict_5min["Wendu"]` shows the height achieved, **in centimeters**, by Wendu in the 5 minute period. Please `round down` your answers to the nearest integer values. Hint - check out the `math.floor()` function.

<!-- grade_id: cell-795209a88bb2fb19, points: 0, grade: false, solution: true -->
```python
# Create a dictionary named climb_dict_5min that uses the player names as the keys
# to store the height they reached after 5 minutes

# climb_dict_5min = ...

### BEGIN SOLUTION

# Here is the actual solution
climb_dict_5min = {}

for name, attribute in candidate_dict.items():
    height = 0    
    if name not in player_list:
        continue
    if (attribute["age"] > 500):
        height += 1500
    for i in range(5):
        if (attribute["height"] < 1200):
            height += attribute["height"] * 1/2
        else:
            height += attribute["height"] * 1/3
    climb_dict_5min[name] = math.floor(height)

### END SOLUTION    
```

<!-- grade_id: cell-94399f58ba746112, points: 2, grade: true, solution: false -->
```python
assert isinstance(climb_dict_5min, dict), "Variable climb_dict_5min should be a dictionary."
assert hashlib.sha256(json.dumps(len(climb_dict_5min)).encode()).hexdigest() == \
'2c624232cdd221771294dfbb310aca000a0df6ac8b66b696d90ef06fdefb64a3', "The number of items in climb_dict_5min is incorrect."
assert hashlib.sha256(json.dumps(''.join(sorted([key for key in climb_dict_5min]))).encode()).hexdigest() == \
'cfaca8aace4a00f2274aabc8062b1c8b0515e5991d0af06560889d7afbe46778', 'Some of the player names in the dictionary are not correct.'
assert hashlib.sha256(json.dumps(''.join(sorted([key+str(climb_dict_5min[key]) for key in climb_dict_5min]))).encode()).hexdigest() == \
'0a82d5146429bb5dfa5fda137b29a36d0f7f02bbf638bc1777505b7fec532506', 'Some of the height values in the dictionary are not correct.'

print(climb_dict_5min)
print("Your solution is correct!")
```

Output:
```
{'Flort': 2041, 'Hambula': 2497, 'Vilawi': 3547, 'Anhaular': 2801, 'Wendu': 2560, 'Baxla': 2187, 'Uluetto': 2270, 'Reomult': 3937}
Your solution is correct!
```

### 3.3) Identify Top Two Players

Consider the results list of the Climb the Wall competition from exercise 3.2 (the `climb_dict_5min` dictionary). What is the order of the players in the results list ordered by the height to which they climbed? Store this answer in the list `climb_result`, with the first element in the list representing the player who climbed the highest, the second element the second highest, and so on.

Additionally, store the names of the two players who finished the highest and the second highest as `first_player` and `second_player`, respectively.

Hint - you may want to use the function `sorted()` (https://stackoverflow.com/questions/20944483/python-3-sort-a-dict-by-its-values)

<!-- grade_id: cell-1c27c3f79e813777, points: 0, grade: false, solution: true -->
```python
# Create a list named climb_result to store the results of the climbing competiton, with the first player 
# being the first item in the list
# Then store the names of the first two players in separate variables

# climb_result = [...]
#first_player = [...]
#second_player = [...]

### BEGIN SOLUTION

# Here is the actual solution
climb_result = sorted(climb_dict_5min, key=climb_dict_5min.get, reverse=True)
first_player = climb_result[0]
second_player = climb_result[1]

### END SOLUTION
```

```python
hashlib.sha256(json.dumps(''.join(climb_result)).encode()).hexdigest()
```

Output:
```
'a012052acac0db65b6af315822cb9dd361375f36bd9b2417e99d52653f6c591f'
```

<!-- grade_id: cell-bb53cd4163fb8b7f, points: 0, grade: true, solution: false -->
```python
assert isinstance(climb_result, list), "climb_result should be a list."
assert hashlib.sha256(json.dumps(len(climb_result)).encode()).hexdigest() == '2c624232cdd221771294dfbb310aca000a0df6ac8b66b696d90ef06fdefb64a3', "The number of items in climb_result is incorrect."
assert hashlib.sha256(json.dumps(''.join(climb_result)).encode()).hexdigest() == 'a012052acac0db65b6af315822cb9dd361375f36bd9b2417e99d52653f6c591f', 'The order of players in climb_result is not correct.'
assert hashlib.sha256(json.dumps(first_player).encode()).hexdigest() == '24029ead2901c23cd16269bb4f5d0448c2f3c3d3f7d10b1bbfa5e35533281137', "The name of first_player is incorrect."
assert hashlib.sha256(json.dumps(second_player).encode()).hexdigest() == '58bd1fe653cf915a620289aceb40938d3a0c88d3338b77f379d64c50b35bece4', "The name of second_player is incorrect."
print(f'First player: {first_player}, second player: {second_player}')
print(climb_result)
print("Your solution is correct!")
```

Output:
```
First player: Reomult, second player: Vilawi
['Reomult', 'Vilawi', 'Anhaular', 'Wendu', 'Hambula', 'Uluetto', 'Baxla', 'Flort']
Your solution is correct!
```

### 3.4) Fill that Jar (Relay) - Optional (ungraded)

The next challenge in Game X is a team competition. Two teams have been formed, with the following players:

    team_1 : ['Reomult', 'Anhaular', 'Wendu', 'Uluetto']
    team_2 : ['Vilawi', 'Flort', 'Hambula', 'Baxla']

They enter a relay-style challenge where each team needs to fill a tank with `1000 litres of water`. The players will take turns in filling the tank.

In each turn, the first player of each team carries water to the tank for one minute, then the second players carry water for one minute, and so on. After each player has had a turn, the first players start again. This process continues until the entire tank has been filled.

It so happens that the amount each player manages to carry at each turn (in one minute) is proportional to their `strength * weight`, i.e. 
     
$$ water\ filled = player\ strength * player\ weight * \frac{1}{50} $$     

Use loops and conditionals to determine which team fills their tank first. Determine also who is the `last player in the winning team` and `how many turns` did it take to fill the tank.

<!-- grade_id: cell-6480d4fd20c9ba1f, points: 0, grade: false, solution: true -->
```python
# Determine which is the winning team, the last player in the winning team who carried water to the tank,
# and how many turns did it take
# To verify that your program is working correctly, it might be helpful to store the player names in each round
# as well as using lists to capture the change of water levels for each team

team_1 = ['Reomult', 'Anhaular', 'Wendu', 'Uluetto']
team_2 = ['Vilawi', 'Flort', 'Hambula', 'Baxla']
    
# winning_team = ...  team_1 or team_2
# last_player = ...
# number_of_turns = ...

### BEGIN SOLUTION

# Here is the actual solution
number_of_turns = 0
tank_1 = 0
tank_2 = 0

tank_1_levels = [0]
tank_2_levels = [0]

player_names1 = []
player_names2 = []

while (tank_1 < 1000 and tank_2 < 1000):
    player_num = number_of_turns % len(team_1)
    
    tank_1 += candidate_dict[team_1[player_num]]["strength"] *  candidate_dict[team_1[player_num]]["weight"] / 50
    tank_2 += candidate_dict[team_2[player_num]]["strength"] *  candidate_dict[team_2[player_num]]["weight"] / 50
    
    tank_1_levels.append(tank_1)
    tank_2_levels.append(tank_2)
    player_names1.append(team_1[player_num])
    player_names2.append(team_2[player_num])        
    
    if tank_1 >= 1000 and tank_2 >= 1000:
        print("tie")
    elif tank_1 >= 1000:
        last_player = team_1[player_num]
        winning_team = "team_1"
    elif tank_2 >= 1000:
        last_player = team_2[player_num]
        winning_team = "team_2"
    
    number_of_turns += 1

### END SOLUTION
```

<!-- grade_id: cell-83b070269298a5c0, points: 0, grade: true, solution: false -->
```python
assert isinstance(winning_team, str), "Review the data type of winning_team."
assert isinstance(last_player, str), "Review the data type of last_player."
assert isinstance(number_of_turns, int), "Review the data type of number_of_fills."

assert hashlib.sha256(json.dumps(winning_team).encode()).hexdigest() == 'c59616eb4cba06a364fa2c79b4d09b3a556d7d6e02c1c0826e06a8c94daa65fc', "winning_team is incorrect."
assert hashlib.sha256(json.dumps(last_player).encode()).hexdigest() == '24029ead2901c23cd16269bb4f5d0448c2f3c3d3f7d10b1bbfa5e35533281137', "last_player is incorrect."
assert hashlib.sha256(json.dumps(number_of_turns).encode()).hexdigest() == 'ef2d127de37b942baad06145e54b0c619a1f22327b2ebbcfbec78f5564afe39d', "The number_of_fills is incorrect."
print(f'winning team: {winning_team}, last player: {last_player}, number of turns: {number_of_turns}')
print("Your solution is correct!")
```

Output:
```
winning team: team_1, last player: Reomult, number of turns: 5
Your solution is correct!
```

## Part 4 - List Comprehension

### 4.1) Identify Digits in Strings

The prize gallery in Game X is a labyrinth with `1000 rooms` numbered from `0 to 999`. The playmaker informs the team that the prizes are stored in a `room with room number divisible by 3`.

Use List Comprehension to save a list of rooms potentially containing prizes, in the list `possible_prize_rooms`. The list should contain the rooms in ascending order.

Hint - if you are new to list comprehension, it may help to write out how you would implement the solution using a traditional for loop first. It could be easier to convert it to the list comprehension after you nailed down the logic.

<!-- grade_id: cell-2fbc492c5b843c6d, points: 0, grade: false, solution: true -->
```python
# Use list comprehension to identify rooms potentially containing prizes

# possible_prize_rooms = [...]

### BEGIN SOLUTION

# Here is the actual solution
possible_prize_rooms = [i for i in range(1000) if i % 3 == 0]

### END SOLUTION
```

<!-- grade_id: cell-3db9af8dbeef62e4, points: 2, grade: true, solution: false -->
```python
assert isinstance(possible_prize_rooms, list), "Variable possible_prize_rooms should be a list."
assert hashlib.sha256(json.dumps(len(possible_prize_rooms)).encode()).hexdigest() == '058d5d43bf485bf78dda1ed4eaf8b78e3106f3c6364c625ead2cc3aeb1908237', "The number of items in possible_prize_rooms is incorrect."
assert hashlib.sha256(json.dumps(sum(possible_prize_rooms)).encode()).hexdigest() == 'a4da607160e448dffd2544a541c229b12825542c08b402b340ba9f102cfc90e2', 'The possible_prize_room list is not correct.'
assert hashlib.sha256(json.dumps(''.join([str(i) for i in possible_prize_rooms])).encode()).hexdigest() == '9cd6260f7e975ea0c5796979925e5f709e3df84040bcd31a1fde25ffb04f55cf', 'The order of rooms in the list is not correct.'
print(possible_prize_rooms)
print("Your solution is correct!")
```

Output:
```
[0, 3, 6, 9, 12, 15, 18, 21, 24, 27, 30, 33, 36, 39, 42, 45, 48, 51, 54, 57, 60, 63, 66, 69, 72, 75, 78, 81, 84, 87, 90, 93, 96, 99, 102, 105, 108, 111, 114, 117, 120, 123, 126, 129, 132, 135, 138, 141, 144, 147, 150, 153, 156, 159, 162, 165, 168, 171, 174, 177, 180, 183, 186, 189, 192, 195, 198, 201, 204, 207, 210, 213, 216, 219, 222, 225, 228, 231, 234, 237, 240, 243, 246, 249, 252, 255, 258, 261, 264, 267, 270, 273, 276, 279, 282, 285, 288, 291, 294, 297, 300, 303, 306, 309, 312, 315, 318, 321, 324, 327, 330, 333, 336, 339, 342, 345, 348, 351, 354, 357, 360, 363, 366, 369, 372, 375, 378, 381, 384, 387, 390, 393, 396, 399, 402, 405, 408, 411, 414, 417, 420, 423, 426, 429, 432, 435, 438, 441, 444, 447, 450, 453, 456, 459, 462, 465, 468, 471, 474, 477, 480, 483, 486, 489, 492, 495, 498, 501, 504, 507, 510, 513, 516, 519, 522, 525, 528, 531, 534, 537, 540, 543, 546, 549, 552, 555, 558, 561, 564, 567, 570, 573, 576, 579, 582, 585, 588, 591, 594, 597, 600, 603, 606, 609, 612, 615, 618, 62
```

### 4.2) Challenge Question on List Comprehension - optional (ungraded)

Some rooms in the prize gallery contain nasty bugs and explosives, and you would want to avoid them. The playmaker kindly informs the team that these rooms are divisible by at least one digit from 4 to 9. Use nested list comprehension to find the list of buggy rooms and store them in `possible_buggy_rooms`

<!-- grade_id: cell-57383b2e4efaea1d, points: 0, grade: false, solution: true -->
```python
# Use list comprehension to identify potentially buggy rooms

# possible_buggy_rooms = [...]

### BEGIN SOLUTION

# Here is the actual solution
possible_buggy_rooms = [i for i in range(1000) if any(i % j == 0 for j in range(4,10))]

### END SOLUTION
```

<!-- grade_id: cell-0313ba2b00efff91, points: 0, grade: true, solution: false -->
```python
assert isinstance(possible_buggy_rooms, list), "Variable possible_buggy_rooms should be a list."
assert hashlib.sha256(json.dumps(len(possible_buggy_rooms)).encode()).hexdigest() == '4299da7466df09516d290f7a99c8b7a2fa94766eb94a61c24e1ce8f6ca80af44', "The number of items in possible_buggy_rooms is incorrect."
assert hashlib.sha256(json.dumps(sum(possible_buggy_rooms)).encode()).hexdigest() == '9860269eaf60fb02d8743ff7c4b9499de72ecf19f7389aa5014e76a6d05d6ed6', 'Some values in possible_buggy_rooms are incorrect.'
print(possible_buggy_rooms)
print("Your solution is correct!")
```

Output:
```
[0, 4, 5, 6, 7, 8, 9, 10, 12, 14, 15, 16, 18, 20, 21, 24, 25, 27, 28, 30, 32, 35, 36, 40, 42, 44, 45, 48, 49, 50, 52, 54, 55, 56, 60, 63, 64, 65, 66, 68, 70, 72, 75, 76, 77, 78, 80, 81, 84, 85, 88, 90, 91, 92, 95, 96, 98, 99, 100, 102, 104, 105, 108, 110, 112, 114, 115, 116, 117, 119, 120, 124, 125, 126, 128, 130, 132, 133, 135, 136, 138, 140, 144, 145, 147, 148, 150, 152, 153, 154, 155, 156, 160, 161, 162, 164, 165, 168, 170, 171, 172, 174, 175, 176, 180, 182, 184, 185, 186, 188, 189, 190, 192, 195, 196, 198, 200, 203, 204, 205, 207, 208, 210, 212, 215, 216, 217, 220, 222, 224, 225, 228, 230, 231, 232, 234, 235, 236, 238, 240, 243, 244, 245, 246, 248, 250, 252, 255, 256, 258, 259, 260, 261, 264, 265, 266, 268, 270, 272, 273, 275, 276, 279, 280, 282, 284, 285, 287, 288, 290, 292, 294, 295, 296, 297, 300, 301, 304, 305, 306, 308, 310, 312, 315, 316, 318, 320, 322, 324, 325, 328, 329, 330, 332, 333, 335, 336, 340, 342, 343, 344, 345, 348, 350, 351, 352, 354, 355, 356, 357, 360, 364, 365,
```

### 4.3) Symmetric matrix

As a final challenge, the playmaker asks the players to create the following list of lists (matrix) using List Comprehension. 

$
\begin{bmatrix}
\begin{bmatrix}0 & 1 & 2 & 3 & 4\end{bmatrix},\\
\begin{bmatrix}1 & 2 & 3 & 4 & 5\end{bmatrix},\\
\begin{bmatrix}2 & 3 & 4 & 5 & 6\end{bmatrix},\\
\begin{bmatrix}3 & 4 & 5 & 6 & 7\end{bmatrix},\\
\begin{bmatrix}4 & 5 & 6 & 7 & 8\end{bmatrix}\space\\
\end{bmatrix}
$

<!-- grade_id: cell-36f6524e13880590, points: 0, grade: false, solution: true -->
```python
# Create the list of lists above using list comprehension.
# The result should be a list assigned to matrix_list. 

# matrix_list = ...

### BEGIN SOLUTION

# Here is the actual solution
matrix_list = [[i + j for i in range(5)] for j in range(5)]

### END SOLUTION
```

<!-- grade_id: cell-989fbfb752329e10, points: 2, grade: true, solution: false -->
```python
assert isinstance(matrix_list, list), "Review the data type of matrix_list."
assert len(matrix_list) == 5, "The length of matrix_list is incorrect."
assert sum([len(l)==5 for l in matrix_list])==5, "The length of some elements of matrix_list is incorrect."
assert sum([sum(i) for i in matrix_list]) == 100, "Some values are incorrect."
print(matrix_list)
print("Your solution is correct!")
```

Output:
```
[[0, 1, 2, 3, 4], [1, 2, 3, 4, 5], [2, 3, 4, 5, 6], [3, 4, 5, 6, 7], [4, 5, 6, 7, 8]]
Your solution is correct!
```

### Extra fact:

There is a [story](http://mathandmultimedia.com/2010/09/15/sum-first-n-positive-integers/) about how [Gauss](https://en.wikipedia.org/wiki/Carl_Friedrich_Gauss#Anecdotes) discovered a formula to calculate the sum of the first `n` positive integers. If there is a formula or a simple calculation that you can perform to avoid extensive code, why not use it instead? 

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2024/blob/main/weekly-workflow.md#link-to-grading)