# SLU05 - Python flow control - Practice


### Foreword for beginner students coming from the Exercise notebook 0
This exercise notebook is step up in that it contains a story. So you'll first have to translate the problem from story into the language of variables, conditions, and loops. It's like when you started learning mathematics in school. First, you learned `4 - 1 - 2 = 1`, then you learned solving problems like `Maria bought four apples, ate one, and gave two to her sister. How many apples has Maria now?`

After you read the problem statement, first identify what are the variables that you need and what are their initial and final values. Then think about the steps you will need to execute to get from the initial values to the result and write them down in plain English. Then, think about the tools you'll need. Will you need an `if ` statement to decide between two possible outcomes? What will be the condition? Will you need a `for` loop to repeat something a defined number of times? Will you need a `while` loop to repeat something an indefinite number of times? Will you need one or more loops or conditions? Finally, think about how to test if your code is working correctly. Can you print out the values of your variables along the way to see what happens?

Write the code in small chunks and always test if it does what you want.

### Welcome to Planet X

Your spaceship has only just landed when you learn that on Planet X only beings with a minimum knowledge of flow control in Python are allowed to leave the station. 
You are guided to a tunnel that will either be the entry to the wondrous world of Planet X where you will be admitted as a spectator of the Games X or ... your end. 

**Hint:** Should you by any accident or intention create an infinite loop (you were duly warned about those in the Learning Notebooks) try interrupting the Kernel, for example by clicking on Kernel in the menu, then on Interrupt!

```python
# Just for evaluating the results.
import math
import json
import hashlib
```

### Exercise 1

Upon entering the tunnel you see a tile in the floor reading: *Which of the following statements will be True? Step on the correct answer to cross this tile!*

A:  3 < 2

B:  A==a

C:  3!=3

D:  23<=23

Which is the answer to step on?

<!-- grade_id: cell-23a99139ad0a0a5a, points: 0, grade: false, solution: true -->
```python
# Your answer here:
# example:
# answer_1 = "E"

### BEGIN SOLUTION
answer_1 = "D"
### END SOLUTION
```

<!-- grade_id: cell-acb79987f258e801, points: 2, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer_1.strip()).encode()).hexdigest() == \
'8f097023401b3d58704c1f3abea8ec2eba57c387f09d10125cd66076afd13b1e',\
"Wrong answer, you fall through the floor tile into an abyss!"
print("Congratulations, you passed the tile!")
```

Output:
```
Congratulations, you passed the tile!
```

### Exercise 2

You continue down the corridor. A screen appears in front of you and asks you to translate the following text into python code:
- create two variables, a variable called "a" should be the integer 1 and a variable called "b" should be the integer ninety-nine
- create a variable called "expression1", a boolean value, with value "a is larger than b"
- create a variable called "expression2", a boolean value, with value "a is smaller than b"
- create a variable called "expression3", a boolean value, with value "at least one of the expressions expression1 or expression2 is true"
- create a variable called "expression4", a boolean value, with value "both of the expressions expression1 and expression2 are true"


<!-- grade_id: cell-66f5c4990f9ec1f4, points: 0, grade: false, solution: true -->
```python
#Your answer here:
#a=
#b=
#expression1=
#expression2=
#expression3=
#expression4=

### BEGIN SOLUTION

a = 1
b = 99
expression1 = False
expression2 = True
expression3 = True
expression4 = False

### END SOLUTION
```

<!-- grade_id: cell-60d3fc8fad22ce3a, points: 3, grade: true, solution: false -->
```python
assert isinstance(a,int) and isinstance(b,int), 'The variables a and b should be integers.'
assert hashlib.sha256(json.dumps(str(a)).encode()).hexdigest() == \
'391552c099c101b131feaf24c5795a6a15bc8ec82015424e0d2b4274a369a0bf',\
"The values of variables a or b are wrong. You hear a small noise and everything around you becomes black!"
assert isinstance(expression1,bool) and isinstance(expression2,bool) and isinstance(expression3,bool) \
and isinstance(expression4,bool), 'The expressions should be booleans.'
assert  hashlib.sha256(json.dumps(''.join([str(expression1),str(expression2),str(expression3),str(expression4)])).encode()).hexdigest()\
== '94514b1f03f079eed63af701925a27f1526ea8bd1f58032e1f6e1fdfb1783944', \
"The boolean values of at least one of the expressions is not correct. You hear a small noise and everything around you becomes black!"
print("The screen vanishes, you may pass!")
```

Output:
```
The screen vanishes, you may pass!
```

### Exercise 3

Just a few more steps down the corridor a robot rolls in your way and asks:

*What text will the following code print?*

<img src="./media/easy_if.PNG" />

<!-- grade_id: cell-937604c83bbcb2a5, points: 0, grade: false, solution: true -->
```python
# Your answer here:
# example:
# answer_3 = "E"

### BEGIN SOLUTION
answer_3 = "B"
### END SOLUTION
```

<!-- grade_id: cell-12dd76da7188130c, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer_3.strip()).encode()).hexdigest() == \
'955cca1ceba45052d85984d3a2565f4ce25b7488602c60a165598bf80b26e472',\
"The robot starts screaming 'Exterminate! Exterminate!', it seems it does not like your answer."
print("The robot deactivates and you pass.")
```

Output:
```
The robot deactivates and you pass.
```

### Exercise 4 

You continue down the corridor and around a corner where you find a table with an elderly lady sitting next to it, knitting. She asks you to use a loop to create a list called `var` with every other lowercase character, like this: `['a', 'c', 'e', ...]`

When you look at her in desperation, she mumbles the words `ord` and `chr`.

<!-- grade_id: cell-6517ffff98b0df00, points: 0, grade: false, solution: true -->
```python
# Hint: use the functions ord and chr you from learning notebook 1 to transform characters
# to code point values that you can use in a loop
#var = []

### BEGIN SOLUTION

var = []
for i in range(ord('a'),ord('z'),2):
    var.append(chr(i))

### END SOLUTION
```

<!-- grade_id: cell-c5b8e750bbb157be, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(''.join(var)).encode()).hexdigest() == \
'fae9e4abc7f52bca1225d60989498b722e9af840996c276fa4a90b629c7f45da',\
"The lady looks at you, disapproval in her eyes. You realize you cannot move nor talk anymore."
print(var)
print("The lady nods and says 'very well, lets take this up a notch, shall we?'")
```

Output:
```
['a', 'c', 'e', 'g', 'i', 'k', 'm', 'o', 'q', 's', 'u', 'w', 'y']
The lady nods and says 'very well, lets take this up a notch, shall we?'
```

### Exercise 5

She tilts her head to the side... You did not think passing me was that easy, did you? Use a `while` loop to iterate over the string "EMIT DRAWER" and put all its characters in the list `answer_5`, but backwards! Like this: `['R','E',...]`.

<!-- grade_id: cell-51863afcdfd87984, points: 0, grade: false, solution: true -->
```python
# Hint: remember that you can access the characters in the string through indexing,
# like this: string[0], which gives 'E'

string = "EMIT DRAWER"
#answer_5 = []

### BEGIN SOLUTION
answer_5 = []
n = len(string)
while n>0:
    answer_5.append(string[n-1])
    n-=1

### END SOLUTION
```

<!-- grade_id: cell-809e87e4da6e7f5a, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(''.join(answer_5)).encode()).hexdigest() == \
'af6c9818ad0d7818936c997defea43369cb430d161a6c5c6d13400f3bebf6469', \
"The lady snips her fingers and you are being attacked by a gigantic alien kitten!"
print(answer_5)
print("Well done, not much more and you shall be free. While you walk down the corridor you feel the lady's eyes in your back.")
```

Output:
```
['R', 'E', 'W', 'A', 'R', 'D', ' ', 'T', 'I', 'M', 'E']
Well done, not much more and you shall be free. While you walk down the corridor you feel the lady's eyes in your back.
```

### Exercise 6

You arrive at a wall that appears to be the end of the corridor. On the wall you discover an inscription with yet another task:

From the `welcome_string` below extract the characters with the indices given in the list `solution_index`. Assign the extracted characters to the `answer_6` list.

- welcome_string = "You are weirdly most welcome to Planet X!"
- solution_index = [4,8,9,18,25,26,27]

<!-- grade_id: cell-f6f8f34235beb39d, points: 0, grade: false, solution: true -->
```python
welcome_string ="You are weirdly most welcome to Planet X"
solution_index=[4,8,9,18,25,26,27]
# answer_6 = []

### BEGIN SOLUTION

answer_6 = []
for i in solution_index:
    answer_6.append(welcome_string[i])

### END SOLUTION
```

<!-- grade_id: cell-0287971b5a29f7a0, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(''.join(answer_6)).encode()).hexdigest() == \
'588e7cfa7a6b57825f0093dded728af014211c4a1d2bc784051da5cfc7c5416f',\
"A net falls on top of you and you are being pulled in a windowless chamber above the corridor"
print(answer_6)
print("The wall vanishes and you blink as your eyes adjust to the sunshine of Planet X")
```

Output:
```
['a', 'w', 'e', 's', 'o', 'm', 'e']
The wall vanishes and you blink as your eyes adjust to the sunshine of Planet X
```

### Exercise 7

Outside you meet an inhabitant of Planet X. They suggest you could win a special spectator seat for the Games X if you manage to solve this exercise! Sounds great!

From a range of numbers from 1 to 77 find the ones that are divisible without rest by 2 and 3 and 7 and put them in the `answer_7` list.

(You will even get a glass of champagne if you manage to write the solution in one line using list comprehension. Otherwise just coffee. ;))

<!-- grade_id: cell-b7d231b0a70c2bbf, points: 0, grade: false, solution: true -->
```python
#answer_7 = []

### BEGIN SOLUTION

answer_7 = []
for i in range(1,78):
    if i%2==0 and i%3==0 and i%7==0:
        answer_7.append(i)

# one line solution
#answer_7 = [i for i in range(1,78) if i%2==0 and i%3==0 and i%7==0]

### END SOLUTION
```

<!-- grade_id: cell-ac56d8db040172ff, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer_7).encode()).hexdigest() == \
'b86b1ea11b28136fe5224b9d1e3017b7efb68d4fae0b90c4940e0c0f89b3907a', \
"No solution no special seat - they say and vanish into the crowd - fair enough"
print(answer_7)
print("You proudly look to the person to present your correct solution - to find they vanished into thin air.")
```

Output:
```
[42]
You proudly look to the person to present your correct solution - to find they vanished into thin air.
```

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2024/blob/main/weekly-workflow.md#link-to-grading)

# Give us feedback!
We'd like to know how you found the learning experience with this SLU. You can leave anonymous feedback in this [feedback form](https://docs.google.com/forms/d/e/1FAIpQLSebiso8S1bfTkELeTqEC18fRy1c79QssrwCbuOgsrTgxWA19w/viewform?usp=sf_link) and help us improve the unit next year. Thank you!