# SLU07 - Function Fundamentals: Exercise notebook

Let's practice some functions! The important part here is that you can take some plain old english and translate it into a function.

```python
# Don't mind this here, just some imports for the exercises - MAKE SURE YOU RUN THIS CELL
import inspect
from io import StringIO
import sys

class Capturing(list):
    def __enter__(self):
        self._stdout = sys.stdout
        sys.stdout = self._stringio = StringIO()
        return self
    def __exit__(self, *args):
        self.extend(self._stringio.getvalue().splitlines())
        del self._stringio    # free up some memory
        sys.stdout = self._stdout
```

### Exercise 1

Let's start with an easy one!

Implement a function that:
- takes no argument
- returns nothing
- is called "print_hello"
- prints the string "Hello"

If you are wondering what this might be useful for - imagine instead of "Hello" you wanted to print something long and nasty to type ;)

<!-- grade_id: cell-cd932fd681964aea, points: 0, grade: false, solution: true -->
```python
# def print_hello():

### BEGIN SOLUTION
def print_hello():
    print("Hello")
### END SOLUTION
```

<!-- grade_id: cell-8d3d93cda67dee49, points: 1, grade: true, solution: false -->
```python
source = inspect.getsource(print_hello)
sig = inspect.signature(print_hello)

with Capturing() as output:
    retval = print_hello()

assert retval == None
assert output[0] == "Hello"

assert 'print' in source
assert 'return' not in source
assert len(sig.parameters) == 0

print("Well done, a start was made")
```

Output:
```
Well done, a start was made
```

### Exercise 2

Let's start with a simple one!

Implement a function that:
- Is named `fun_adding_numbers`
- Takes two arguments
- Adds the arguments together
- Returns the value of the numbers added together


<!-- grade_id: cell-a3913c46bc1f65f2, points: 0, grade: false, solution: true -->
```python
# def fun_adding_numbers(a, b):

### BEGIN SOLUTION
def fun_adding_numbers(a, b):
    return a + b
### END SOLUTION
```

<!-- grade_id: cell-b7d97d9eb6d71625, points: 2, grade: true, solution: false -->
```python
sig = inspect.signature(fun_adding_numbers)

assert fun_adding_numbers(1, 2) == 3
assert fun_adding_numbers(4, 5) == 9
assert '+' in inspect.getsource(fun_adding_numbers)
assert len(sig.parameters) == 2
print("No mistakes detected, well done!")
```

Output:
```
No mistakes detected, well done!
```

### First function done!! It's as simple as that! ðŸ˜Ž 

### Exercise 3

Let's go for another one! 

Implement a function that

- Is named `fun_multiplying_numbers`
- Takes two parameters
- Multiplies the arguments together
- Returns the value of the numbers multiplied together

No worries about creating functions that do almost the same, it will all be reused later on


<!-- grade_id: cell-b5cba4d5bddbb4b5, points: 0, grade: false, solution: true -->
```python
# def fun_multiplying_numbers(a, b)

### BEGIN SOLUTION
def fun_multiplying_numbers(a, b):
    return a * b
### END SOLUTION
```

<!-- grade_id: cell-5db0d0c172007a90, points: 1.5, grade: true, solution: false -->
```python
source = inspect.signature(fun_multiplying_numbers)

assert fun_multiplying_numbers(1, 2) == 2
assert fun_multiplying_numbers(4, 5) == 20
assert '*' in inspect.getsource(fun_multiplying_numbers)
assert len(source.parameters) == 2
print ("Could not find a problem here! Functions are not all that bad, right?")
```

Output:
```
Could not find a problem here! Functions are not all that bad, right?
```

### Exercise 4

Now for some special functions points! You need to think about this one and what you have learned up until now.

Try to implement the previous `fun_multiplying_numbers` function but now with a twist:
- Is named `fun_multiplying_numbers_with_a_twist`
- Takes two parameters
- **Multiplies the arguments without using the * operator!**
- Returns the value of the multiplication

Tip: If you remember the good old days of math, multiplication is just a sequence of adding operations in a loop and you already learned how to do those!

<!-- grade_id: cell-24f09cf30d150f7a, points: 0, grade: false, solution: true -->
```python
# def fun_multiplying_numbers_with_a_twist(a,b):
    # result = 0
    # (loop)
    # return result

### BEGIN SOLUTION
def fun_multiplying_numbers_with_a_twist(a, b):
    result = 0
    for i in range(0, a):
        result += b
    return result
### END SOLUTION
```

<!-- grade_id: cell-9038267e2c8c00c4, points: 2, grade: true, solution: false -->
```python
source = inspect.signature(fun_multiplying_numbers_with_a_twist)

assert fun_multiplying_numbers_with_a_twist(1, 2) == 2
assert fun_multiplying_numbers_with_a_twist(4, 5) == 20
assert fun_multiplying_numbers_with_a_twist(100, 200) == 20000
assert '*' not in inspect.getsource(fun_multiplying_numbers_with_a_twist), "Don't use the multiplication operator! ðŸ˜‰"
assert '/' not in inspect.getsource(fun_multiplying_numbers_with_a_twist), "Dividing by the inverse is the same as using the multiplication operator! ðŸ˜‰"
assert len(source.parameters) == 2
print ("All asserts passed!")
```

Output:
```
All asserts passed!
```

### Exercise 5

Implement a function that

- Is named `concat_strings`
- Takes two parameters that are strings
- Concatenates the strings together in the order that they are received
- Returns the concatenated strings
- **Twist**: the second parameter assumes as a default string " best" -> *mind the initial space!*

<!-- grade_id: cell-90e04d850de62687, points: 0, grade: false, solution: true -->
```python
#Remember the default values ðŸ¤”
#def concat_strings(...
### BEGIN SOLUTION
def concat_strings(a, b=" best"):
    return a + b
### END SOLUTION
```

<!-- grade_id: cell-d309394f7d502266, points: 2, grade: true, solution: false -->
```python
source = inspect.getsource(concat_strings)

sig = inspect.signature(concat_strings)

assert concat_strings('hello', ' world') == 'hello world'
assert concat_strings('Daniel', ' will give me a good grade ðŸ™') == 'Daniel will give me a good grade ðŸ™'
assert concat_strings('This course is the') == 'This course is the best'
assert concat_strings('functions are the') == 'functions are the best'
assert len(sig.parameters) == 2
print (concat_strings('You are the'))
```

Output:
```
You are the best
```

### Exercise 6

Implement a function that

- Is named `convert_to_int_and_add`
- Takes two parameters
- Both of the parameters are integers **represented as strings** (for example, "9")
- Converts the strings to integers
- Add the integers together
- Return the value of the integers added together

<!-- grade_id: cell-5769f897199017e3, points: 0, grade: false, solution: true -->
```python
#def convert_to_int_and_add(a, b):
### BEGIN SOLUTION
def convert_to_int_and_add(a, b):
    return int(a) + int(b)
### END SOLUTION
```

<!-- grade_id: cell-bf3c5cd42c49fa00, points: 2, grade: true, solution: false -->
```python
source = inspect.getsource(convert_to_int_and_add)

sig = inspect.signature(convert_to_int_and_add)

assert convert_to_int_and_add('1', '2') == 3
assert convert_to_int_and_add('5', '100') == 105
assert '+' in source
assert 'int' in source
assert len(sig.parameters) == 2
print("You are getting the hang of it!")
```

Output:
```
You are getting the hang of it!
```

### Exercise 7

Implement a function that

- Is named `two_step`
- Takes 4 parameters
- Multiplies the first 2 together and stores the output in a variable called `step1`
- Divides 3rd parameter by the 4th parameter and stores the output in a variable called `step2`
- Adds the variables `step1` and `step2` together, calling it `result`
- If `result` is greater than 20, subtract 10 to it
- Returns the final result

<!-- grade_id: cell-70188605f9305dc8, points: 0, grade: false, solution: true -->
```python
#def two_step(a,b,c,d):
### BEGIN SOLUTION
def two_step(a, b, c, d):
    step1 = a * b
    step2 = c / d

    result = step1 + step2
    if result > 20:
        result -= 10

    return result
### END SOLUTION
```

<!-- grade_id: cell-9712fb50a8f3d533, points: 2, grade: true, solution: false -->
```python
source = inspect.getsource(two_step)

sig = inspect.signature(two_step)

assert two_step(2, 2, 4, 2) == 6.0, "The final result is not matching"
assert two_step(2, 3, 100, 5) == 16.0, "Are you subtracting correctly? ðŸ¤”"
assert '*' in source
assert '/' in source
assert 'step1' in source
assert 'step2' in source
assert 'result' in source
assert len(sig.parameters) == 4
print("Good!")
```

Output:
```
Good!
```

### Exercise 8

Write a function that

1. Is named `holla`
1. Takes 5 arguments
1. Calls `fun_adding_numbers` with the first two arguments and stores the output in `out1`
1. Calls `fun_multiplying_numbers` with the second two arguments and stores the output in `out2`
1. Calls `two_step` with the last 4 arguments and stores the output in `out3`
1. Returns the addition of `out1`, `out2`, and `out3`

<!-- grade_id: cell-a5e8cf3762c2538a, points: 0, grade: false, solution: true -->
```python
#def holla(a,b,c,d,e):
### BEGIN SOLUTION
def holla(a, b, c, d, e):
    out1 = fun_adding_numbers(a, b)
    out2 = fun_multiplying_numbers(c, d)
    out3 = two_step(b, c, d, e)
    return out1 + out2 + out3
### END SOLUTION
```

<!-- grade_id: cell-a02c4760a2ce92bf, points: 2.5, grade: true, solution: false -->
```python
source = inspect.getsource(holla)
sig = inspect.signature(holla)

assert holla(1, 2, 2, 4, 2) == 17
assert holla(1, 2, 3, 100, 5) == 319.0
assert 'fun_adding_numbers(' in source
assert 'fun_multiplying_numbers(' in source
assert 'two_step(' in source
assert 'out1+out2+out3' in source.replace(" ","")
assert len(sig.parameters) == 5
print("GREAT!")
```

Output:
```
GREAT!
```

### Exercise 9

Write a function that

1. Is called `add_and_print`
1. Takes two arguments
1. Adds them together
1. Prints the output
1. Returns the value `5`

<!-- grade_id: cell-5a9e41f05baeb9b5, points: 0, grade: false, solution: true -->
```python
#def add_and_print(a,b):
### BEGIN SOLUTION
def add_and_print(a, b):
    print(a + b)
    return 5
### END SOLUTION
```

<!-- grade_id: cell-733a771bc8560db2, points: 1, grade: true, solution: false -->
```python
source = inspect.getsource(add_and_print)
sig = inspect.signature(add_and_print)

with Capturing() as output:
    retval = add_and_print(1, 2)

assert retval == 5
assert output[0] == '3'

with Capturing() as output:
    retval = add_and_print(10287366, 912865763524)

assert retval == 5
assert output[0] == '912876050890'

assert '+' in source
assert len(sig.parameters) == 2
print("Good! and now to your final challenge... hehehehe")
```

Output:
```
Good! and now to your final challenge... hehehehe
```

### Exercise 10a

Write a function that

1. Is called `prepare_list`
1. Takes 1 argument, a list
1. checks if all elements in the lists are real numbers
1. removes the ones that are not real numbers, without changing the order
1. Returns the list with only numbers, if no numbers were in the list, it returns an empty list

<!-- grade_id: cell-db3cb0790727a7b9, points: 0, grade: false, solution: true -->
```python
# l is a list of elements for example [1,2,"a",4, 5.4]
#def prepare_list(l):
### BEGIN SOLUTION
def prepare_list(l):
    return [i for i in l if type(i) in (int,float)]
### END SOLUTION
```

<!-- grade_id: cell-bcba7c8fcfe6a995, points: 2, grade: true, solution: false -->
```python
sig = inspect.signature(prepare_list)
assert len(sig.parameters) == 1
assert prepare_list([1,2,3]) == [1,2,3]
assert prepare_list([5,"a",1,(3,5)]) == [5,1]
assert prepare_list([[1,2,3],"a",3,'test']) == [3]
assert prepare_list([{'k':1,5:'e'}]) == []
print("There you go! almost there...")
```

Output:
```
There you go! almost there
```

### Exercise 10b

And now for your last challenge...  write a function called `list_calc` that upon receiving a list sums up all real numbers in it. Then it multiplies that result by the last list element that is a real number. In case there are no real numbers in the list return 0. Feel free to use what you created in 10a. 

<!-- grade_id: cell-ec81ca06381f7aef, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def list_calc(l):
    cleaned_list=prepare_list(l)
    return_value=0
    if cleaned_list != []:
        sum_value=0
        for i in cleaned_list:
            sum_value+=i
        return_value=sum_value*cleaned_list[-1]
    return return_value

### END SOLUTION
```

<!-- grade_id: cell-49471cc1739e12df, points: 2, grade: true, solution: false -->
```python
sig = inspect.signature(prepare_list)
assert len(sig.parameters) == 1
assert list_calc([1,2,-3]) == 0
assert list_calc([1,2,2]) == 10, "don't forget to add the last element before multiplying by it"
assert list_calc([5,"a",1,(3,5)]) == 6
assert list_calc([[1,2,3],"a",3,'test']) == 9, "if only one real number is in the cell, it is also the last element"
assert list_calc([{'k':1,5:'e'}]) == 0, "return 0 if no values are real numbers"
assert list_calc([0.5,0.2,10,-2]) == -17.4, "floats are also real numbers!"
print("Well done! I am proud of you!")
```

Output:
```
Well done! I am proud of you
```

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2024/blob/main/weekly-workflow.md#link-to-grading)