# SLU04 - Data Structures - Exercise notebook 2

### Start by importing the following packages

```python
#used for evaluation
from IPython.core.getipython import get_ipython
import hashlib
import json
```

In this notebook, we will test the following:
- General operations with data structures
- Mutability and immutability
- Zipping and unpacking data structures

**IMPORTANT:** Some exercises require you to use variables defined in previous cells, so don't forget to run them before you run the actual exercise code!

For example, in Exercise 2.1 you need to run the cell after "Remember our zipper from the Learning notebook? Run the following cell." before executing your solution.

---

## Exercise 1: Data structures and mutability

This exercise covers topics regarding the learning notebooks 1 and 2.

### 1.1. Create a tuple of empty lists

Create a tuple named `my_tuple` which holds 3 empty lists. Solve the exercise regarding the following rules:
- Use a single line of code
- Don't repeat any character except for white spaces. (e.g. if you use the `m` character once, you can't use it again in this solution)

<!-- grade_id: cell-7a019784784bde0a, points: 0, grade: false, solution: true -->
```python
#my_tuple = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3286ebcfe53cb78a, points: 3, grade: true, solution: false -->
```python
def filter_comments(code_lines):
    filtered_lines = []
    for line in code_lines:
        line = line.strip()
        if line and not line.startswith('#'):
            filtered_lines.append(line)
    return filtered_lines

source_code = get_ipython().user_ns["In"][-2]
source_code_lines = source_code.strip().splitlines()
solution_lines = filter_comments(source_code_lines)
number_of_lines = len(solution_lines)
solution_code = solution_lines[0]

assert number_of_lines == 1, "Write your solution in a single code line."
assert len(set(solution_code.replace(" ", ""))) == \
       len(solution_code.replace(" ", "")), "You can't repeat characters!"
assert isinstance(my_tuple, tuple), "Are you sure my_tuple is a tuple?"
assert len(my_tuple) == 3, "The length is not quite right."
assert all(isinstance(i, list) for i in my_tuple), "Are the tuple's items lists?"
assert all(len(i) == 0 for i in my_tuple), "Are the lists empty?"

print("Well done! 🙌")
```

### 1.2. Considering the previous exercise, answer the following questions:

1.2.1. Can you assign a `list` holding 3 tuples to the variable `my_tuple`?

> 💡 **Tip:** If you feel tempted to answer `e)`, go back to 'Learning notebook 2' and read the first paragraph under 5.2 ten consecutive times for self-punishment!

a) No, because a tuple (which is immutable) has already been assigned to that variable.  
b) No, because `"tuple"` in the variable's name (`my_tuple`) would collide with the type `list` of the assigned list.  
c) No, because a `list` (which is mutable) can't contain immutable types (like tuples).  
d) Yes, because it would just create a new variable with the same name `my_tuple`, and assign it a list holding 3 tuples.  
e) Yes, because variables are mutable.  
f) Yes, because lists are mutable.

<!-- grade_id: cell-82b80f38223c4dbd, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
#answer = "e"
#answer = "f"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a75826cfdf84d0c1, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633', "Wrong answer."

print("Well done! 🙌")
```

1.2.2. Can you turn the empty lists inside `my_tuple` into empty dictionaries?

a) No, because tuples are immutable  
b) No, because they're both empty  
c) Yes, because they're both empty  
d) Yes, because lists are mutable   
e) Yes, because dictionaries are mutable

<!-- grade_id: cell-cdee69643e449b12, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
#answer = "e"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a9eac81b5d588456, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == 'ac8d8342bbb2362d13f0a559a3621bb407011368895164b628a54f7fc33fc43c', "Wrong answer."

print("Well done! 🙌")
```

1.2.3. Can you insert the following sets pair `{1, 2}`, `{3, 4}` inside each list in `my_tuple`?

a) No, because tuples are immutable  
b) No, because the list is initially empty  
c) Yes, because the list is initially empty  
d) Yes, because lists are mutable   
e) Yes, because sets are mutable

<!-- grade_id: cell-f298899f40c35c56, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
#answer = "e"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1e6a8577155d82b, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633', "Wrong answer."

print("Well done! 🙌")
```

1.2.4. Run the following cell:

```python
a = 1
b = 2
c = 3
d = 4

my_tuple = ({a, b}, {c, d})
my_tuple
```

Now, suppose you assign a dictionary to the variable `a`:

```python
a = {'key': 1}
```

What happens if you now run `my_tuple`?

a) You get an error because a dictionary was assigned to the variable `a` and sets can't hold dictionaries since they're not hashable  
b) You get an error because tuples are immutable  
c) Nothing happens to `my_tuple` because it holds a reference to the set `{1, 2}`, and another reference to the set `{3, 4}` both unrelated to the reference `a` points to  
d) Nothing happens to `my_tuple` because `a` isn't reassigned to the dictionary since it is inside a tuple, (which is immutable)

<!-- grade_id: cell-f32cd95c431008df, points: 0, grade: false, solution: true -->
```python
# Uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f9b55e1465b76e45, points: 3, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '879923da020d1533f4d8e921ea7bac61e8ba41d3c89d17a4d14e3a89c6780d5d', "Wrong answer."

print("Well done! 🙌")
```

---

## Exercise 2: `zip` function and the 'splat' `*` operator

###  2.1. Zip and unzip

Remember our zipper from the Learning notebook? Run the following cell.

```python
teeth_collection = [(1, 3, 5, 7, 9), (2, 4, 6, 8, 10)]
```

Write the expression to zip and unzip the variable `teeth_collection` as a list and assign it to the variable `zip_and_unzip`.  
Similarly to exercise 1.1., use a single line of code to write your solution.

> 💡 **Tip:** Note that `teeth collection` is now a list holding both left and right teeth collections, which means that you probably need to start by unp... the left and right teeth collections.  
> 💡 **Tip:** Start building the expression from the inside out, following the steps outlined in the Learning notebook.  
> 💡 **Tip:** You will "probably" need more the one sequential `zip` function and 'splat' `*` operator combinations for this one.  
> 💡 **Tip:** Read the tips 😉

<!-- grade_id: cell-8a922c6f17dd7c19, points: 0, grade: false, solution: true -->
```python
#zip_and_unzip = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5f15b89580676c8c, points: 5, grade: true, solution: false -->
```python
def filter_comments(code_lines):
    filtered_lines = []
    for line in code_lines:
        line = line.strip()
        if line and not line.startswith('#'):
            filtered_lines.append(line)
    return filtered_lines

source_code = get_ipython().user_ns["In"][-2]
source_code_lines = source_code.strip().splitlines()
solution_lines = filter_comments(source_code_lines)
number_of_lines = len(solution_lines)
solution_code = solution_lines[0]

zips = solution_code.count("zip")
splats = solution_code.count("*")

assert number_of_lines == 1, "Write your solution in a single code line."
solution_code = str(solution_code)
assert zips == 4, "Check the number of times you need to use the `zip` function."
assert splats == 2, "Check the number of times you need to use the 'splat' `*` operator."
assert isinstance(zip_and_unzip, list), "Are you sure you used a list?"
assert len(zip_and_unzip) == 2, "The length is not quite right."
assert zip_and_unzip == teeth_collection, "The final 'zipper' configuration must be exactly like the starting one"

print("Great job!! ✨✨✨✨✨ 🙌\nEnjoy the rest of the warrior. 😴😁")
```

# Submit your work!

To submit your work, [follow the steps in the section "Grading the Exercise Notebook"!](https://github.com/LDSSA/ds-prep-course-2024#22---working-on-the-learning-units)