# SLU14 - Modules and Packages 

Now is the time to evaluate your knowledge on Modules & Packages in Python!\
To answer this quiz, please follow along the instructions provided. If you see: 
```python 
raise NotImplementedError()
```
Delete that line before you test your answer. Otherwise, it will raise an error as it believes that you have not solved the problem yet.

__Tip:__ Once finished with the notebook, it might be a good idea to restart the Kernel and running all cells by selecting `Kernel >> Restart Kernel and Run All Cells...`. This will ensure that all cells run properly before submitting.

#### Initialise Autoreload - Notebook specific topic
Specifically for Exercise 5, you will be asked to add a module to a package/library. Unfortunately, the way Jupyter notebooks work is not very good for development and testing of modules since its *Kernel* is always running.

When importing a module/package, Python will add it to its module dictionary, which you can view with the following code:
```python
import sys
print(sys.modules)
```
When you run an import statement, Python will check if it has a cached value in this dictionary. Having the same *Kernel* running during the execution of a notebook means that cells will not re-import modules/packages that have already been run because they are cached.

By cached what we mean is that changes made to modules that have been imported will not be brought into the current *Kernel*, i.e. you will have an outdated version of the import not matter how many times you re-import it. We will __try__ (this has known quirks and downsides) to combat this by using the `autoreload` jupyter extension (which came pre-built with your jupyter installation). Its job is to force a reload of imported packages, thus hopefully keeping your modules updated.

```python
%load_ext autoreload 
%autoreload 2
```

#### Start by importing the following module

```python
#used for evaluation
import utils
from pathlib import Path
```

### Exercise 1: What is the difference between a class and a module?

a) There is no difference.

b) A module is a set of files with the `.py` extension and a class is just one of those files.

c) A module is a file with the `.py` extension which can contain classes or one class, but can also contains functions and variables.

d) A class is a file with the `.ipynb` extension and a module with the `.py` extension.

<!-- grade_id: cell-ec299e14ecce1f27, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-60bb3a2487965c64, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q1, "Wrong answer."
```

### Exercise 2: How can we run the module as a script?

a) Just create your module with the desired attributes only: classes, functions and/or variables.

b) With no special organization, but with all the code inside `if __init__:` block.

c) With no special organization at all, because it automatically already runs as a script.

d) We only have to make sure that the entry point (the part of the code where the program starts executing) should be inside  `if __name__=='__main__':` block.

<!-- grade_id: cell-2f1a5dfccc2eee9b, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-18368bb83c4a41b4, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q2, "Wrong answer."
```

### Exercise 3: If you create an empty new module (.py file) and print __name__. What will happen and why? 

a) It will give an error, because the module cannot be empty.

b) It will print ```__main__ ```, because python always sets the variable 'name' to 'main' before running the file.

c) It will print ```__name__ ```, because that is exactly what you asked python to do.

<!-- grade_id: cell-7a2c49000d940397, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cd2427e89b7707cb, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q3, "Wrong answer."
```

### Exercise 4: So, now you made your first module and import it in your second module. Why are you happy that you used `if __name__=="__main__"` in your first module?

a) It looks interesting in your first module to use it, but doesn't do anything for your second module.

b) It is not possible to import the first module without the `if __name__=='__main__':` statement.

c) Python checked, during the import, if the name is "main". This is true (as seen in the previous question) and therefor it will only use the attributes within the if-statement and not everything in the first module file.

d)Python checked, during the import, if the name is "main". This is not true, since the name will be set as the name of the imported module. But because of the statement `if __name__=='__main__':`, python knows to set the name back to "main". Now it is clear that the second module is your main file.

<!-- grade_id: cell-a4658c8feaa4f078, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9a98124ab22debe2, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q4, "Wrong answer."
```

### Exercise 5: Inside the directory stores, create a module named bookstore with the following attributes:


- A __class__ named `Book` with three attributes, `title`, `author`, and `price`.
  - *Attributes*: 
    - `title` attribute should be a string representing the title of the book
    - `author` attribute should be a string representing the author of the book
    - `price` attribute should be an integer representing the price of the book
  - *Methods*:
    - `get_book_info()` method should return a string that contains the book's title, author, and price in the following format:
      
      `"Title: [title], Author: [author], Price: [price]."` --> replacing the \[...\] for the actual attribute of the class

      example: `"Title: Don Quixote, Author: Miguel de Cervantes, Price: 8."`
- A __function__ named `get_total_price` that accepts a list of `Book` objects as an input and returns the total price of all the books in the list.
- A __variable__ named `description` with the following string value: `"This is a module named bookstore."`
- 
Now let's see if your module works and follow the steps below to see the total price of your new books. 

To evaluate your module, create a boolean variable named `done` with the value True.

__NOTE:__ As stated in the beginning of the notebook, working with notebooks __while__ developing modules can be a pain. The `autoreload` extension will try to reload your modules in case you need to do some developments to them wile running the notebook. If for some reason this does not work, please try to __restart__ your Kernel by going to `Kernel >> Restart Kernel and Run All Cells...` (or similar option) at the top.

<!-- grade_id: cell-aea5631b25c21caa, points: 0, grade: false, solution: true -->
```python

book_list = [
    ("To Kill a Mockingbird", "Harper Lee", 2),
    ("The Great Gatsby", "F. Scott Fitzgerald", 5),
    ("Pride and Prejudice", "Jane Austen", 7),
    ("The Catcher in the Rye", "J.D. Salinger", 6),
    ("1984", "George Orwell", 15)
    ]
    

# Set done = True

# import your module

# create the list of Book objects
 
# calculate the total price of the books using the get_total_price function

# total_price =

# Print the total price

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-421c00edf62020d8, points: 5, grade: true, solution: false -->
```python
if done:
    try:
        from stores import bookstore
        
        assert total_price == 35, "The total price for the test book list should be 35 euro's"
        assert bookstore, "Did you store a module named bookstore inside stores folder?"
        
        book1 = bookstore.Book("Hamlet", "William Shakespeare", 10)
        assert book1.title == "Hamlet", "Did you set title as an attribute of class Book?"
        assert book1.author == "William Shakespeare", "Did you set auteur as an attribute of class Book?"
        assert book1.get_book_info() == "Title: Hamlet, Author: William Shakespeare, Price: 10.", "Did you set get_book_info correctly as a method of class Book?"

        book2 = bookstore.Book("The Great Gatsby", "F. Scott Fitzgerald", 15)
        assert bookstore.get_total_price, "Did you create a function named get_total_price()?"
        assert bookstore.get_total_price([book1, book2]) == 25, "The total price for the test book list should be 25 euro's"

        assert bookstore.description == "This is a module named bookstore.", "Did you set description well?"
        
        print("Great job! You are awesome :-).")
        
    except ImportError as exc:
        raise ImportError("In the directory 'stores' is no module called 'bookstore'. Try again or ask for help")
else:
    raise NotImplementedError("You have to create a boolean variable, called 'done', with the value True")
```

### Exercise 6: We love a good python course, right?! Import the class Course, from module course, and create a new object from that.

The `library` folder is in the current directory. Explore its contents and understand its structure.

Import the class `Course`, that belongs to module course. You should do this in a way that imports `Course` directly, meaning that you use the syntax 
```python 
course = Course(x,y,z)
```

Afterwards, create a course variable called `this_course`, with the following attributes:
* name: prep-course
* author: ldssa
* year: 2024
  
Look at the class' `__init__` method if you have doubts regarding its initialisation.

<!-- grade_id: cell-8dafeb32a322444f, points: 0, grade: false, solution: true -->
```python
# from module1.module2.etc... import Class,etc..
#this_course = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ac799d1b55a4cba, points: 4, grade: true, solution: false -->
```python
assert Course, "Course class was not imported with success"
assert isinstance(this_course, Course), "this_course is not of type Course"
assert this_course.name == "prep-course", "this_course has the wrong name"
assert this_course.author == "ldssa", "this_course has the wrong author"
assert this_course.year == 2024, "this_course has the wrong year"
```

### Exercise 7: What is a package in Python?

a) A single Python file (with a .py extension) is considered a package.

b) A set of modules that share the same directory and this directory must include an `__init__.py` file.

c) The thing the postman delivers when you order something online.

d)  A set of modules that share the same directory and this directory should not include an `__init__.py` file.

<!-- grade_id: cell-17731565ade4c8d8, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed761f5a351a7094, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q8, "Wrong answer."
```

### Exercise 8: Building a calculator package

Build a python package called `calculator` with the following directory structure:

```
calculator/
├── main.py
└── operations
    ├── 
    ├── addition.py
    └── subtraction.py

```

`addition.py` should define the following function:

```python
def add(a, b):
    return a + b
```

`subtraction.py` should define the following function:

```python
def subtract(a, b):
    return a - b
```

`main.py` should define the following statements:

```python

print(add(2, 3))
print(subtract(5, 1))
```

You should then be able to run the program in a terminal and get the output:

``` python
> python calculator/main.py
5
4

```

Beware:
* the directory diagram above may be missing something... 
* the main file may be missing something...

Were you able to complete this exercise? (if not, ask for help!)

<!-- grade_id: cell-8df2491f6b0042a0, points: 0, grade: false, solution: true -->
```python
answer = None
answer = !python calculator/main.py

# delete the "raise NotImplemented()" below

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-839cdbbdbfebce43, points: 2, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.ops, "Your script did not return the correct response."
```

### Exercise 9: Why are packages important to use for any Python user?

a) Packages in Python were created to make programming more difficult and confusing for beginners.

b) It is not really necessary anymore, nor is it important, but it is just the way it has always been done.

c) Packages in Python are important because they allow you to hide your code from other developers and prevent them from using it.

d) They are used to organize and distribute modules and we love organized. A lot of time you import modules that contains other modules. And this way it is easier to manage and organize your code.

<!-- grade_id: cell-471e5327dcc8e7b7, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9d78f1dc248bb974, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q10, "Wrong answer."
```

### Exercise 10:  What is the command used to install packages in a python environment?

a) cat   
b) pip   
c) package   
d) venv   

<!-- grade_id: cell-6d4d775081c8be95, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-342a8848d6c5e6fe, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q11, "Wrong answer."
```

### Exercise 11: Given the following directory structure, what is the RELATIVE import of the function *hello* (defined in module5.py) considering that you want to import it in module1.py?

```
project_dir
└── package
    ├── __init__.py
    ├── subpackage1
    |   ├── __init__.py
    │   ├── module1.py    --> needs to import hello()
    │   └── module2.py
    └── subpackage2
        ├── __init__.py
        ├── module3.py
        ├── module4.py
        └── subsubpackage1
            ├── __init__.py
            └── module5.py --> defines function hello()
```

*HINT*: Remember that there are two ways to start a relative import, one that starts looking in the current package, and another that will start at the parent package. If you want to work the problem hands-on, use `python -m package.subpackage1.module1` from within `project_dir`.

a) `.subpackage2.subsubpackage1.module5 import hello`

b) `..subsubpackage1.module5 import hello`

c) `package.subpackage2.subsubpackage1.module5 import hello`

d) `..subpackage2.subsubpackage1.module5`

<!-- grade_id: cell-74261e7c00f33fda, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6f894fd89ca8edfb, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q12, "Wrong answer."
```

### Exercise 12: The package NumPy is commonly used in Data Science projects to perform mathematical operations, along with being a common subdependency of many popular packages. It is usually imported with the alias `np`. How can you achieve this?

a) `from numpy import np`

b) `import numpy as np`

c) `import np from numpy`

d) `as np import numpy`

<!-- grade_id: cell-ee696c0c592ba2de, points: 0, grade: false, solution: true -->
```python
# uncomment the right answer
# answer = "a"
# answer = "b"
# answer = "c"
# answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-061154af10aef5ee, points: 1, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q13, "Wrong answer."
```

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2024/blob/main/weekly-workflow.md#link-to-grading)